The TOR signaling pathway regulates vegetative development, aflatoxin biosynthesis, and pathogenicity in Aspergillus flavus

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public Review):

While I acknowledge the authors' effort in conducting Southern blot analysis to address my prior concern regarding the presence of dual copies of torA and tapA, I find their current resolution inadequate. Specifically, the simple deletion of the respective result sections for torA and tapA significantly impacts the overall significance of this study. The repeated unsuccessful attempts to generate correct mutants only offer circumstantial evidence, as technical issues may have been a contributing factor. Therefore, instead of merely removing these sections, it is essential for the authors to present more compelling experimental data demonstrating that torA and tapA are indeed vital for the viability of A. flavus. Such data would enhance the overall significance of this study.

We agree and appreciate reviewer's important comments on our manuscript. In this version, we address this issue by providing additional experimental data to further support the importance of torA and tapA in the viability of A. flavus. We conducted additional experiments to generate more compelling evidence regarding the essential role of torA and tapA in the growth and development of A. flavus. We constructed a mutant strain (xylPtorA) using an xylose-inducible promoter, which allows for conditional induction with the addition of xylose (Lines 204-238, page 10).

Due to the unsuccessful construction of TapA knockout strains and xylose promoter replacement strains, we used homologous recombination to replace the original promoter with the gpdA strong promoter for overexpression of tapA (OE::tapA). We thank reviewer for highlighting this important aspect, and we revise our manuscript accordingly to enhance its overall significance (Lines 277-297, page 13). We are grateful for the opportunity to enhance our manuscript and believe these revisions provide a more comprehensive understanding of the roles of torA and tapA in A. flavus.

Reviewer #1 (Recommendations For The Authors):

Minor comments

Lines 421-423 and 465-466: these sentences are grammatically awkward. Please rephrase them.

Thank you for your feedback on our manuscript. We conducted additional experiments, so we have removed the sentence from the manuscript to maintain coherence and avoid redundancy.

Reviewer #2 (Public Review):

In this study, authors identified TOR, HOG and CWI signaling network genes as modulators of the development, aflatoxin biosynthesis and pathogenicity of A. flavus by gene deletions combined with phenotypic observation. They also analyzed the specific regulatory process and proposed that the TOR signaling pathway interacts with other signaling pathways (MAPK, CWI, calcineurin-CrzA pathway) to regulate the responses to various environmental stresses. Notably, they found that FKBP3 is involved in sclerotia and aflatoxin biosynthesis and rapamycin resistance in A. flavus, especially that the conserved site K19 of FKBP3 plays a key role in regulating aflatoxin biosynthesis. In general, the study involved a heavy workload and the findings are potentially interesting and important for understanding or controlling the aflatoxin biosynthesis. However, the findings have not been deeply explored and the conclusions mostly are based on parallel phenotypic observations.

Thank you for your constructive comments on our manuscript. In response to your comments, we have conducted additional experiments, including the construction of a xylose promoter mutant strain and an overexpression strain. We have also expanded the discussion section to provide a more comprehensive analysis of our findings in the context of existing literature. Thank you again for your insightful feedback, which has been instrumental in improving the quality of our work. (Lines 464-469, page 22).

Reviewer #2 (Recommendations For The Authors):

Point 1: Our findings revealed that both the tor and tapA genes are present in double copies in our strains, which guided our decision to construct single-copy deletion strains using homologous recombination However, the tor gene in A. flavus exhibited varying copy numbers, as was confirmed by absolute quantification PCR at the genome level (Table S1). However, it is hard to understand for Table S1: Estimation of copy number of tor gene in A. flavus toro and sumoo stand for the initial copy number, and the data are graphed as the mean {plus minus} 95%confidence limit. CN is copy number. As indicated in the Methods, Using sumo gene as reference, the tor and tapA gene copy number was calculated by standard curve. In Table S1 of WT, for tor gene, CN value is1412537 compared to 1698243 in tor+/-, for the reference gene sumo,794328 compared to1584893, how these data could support copy gene numbers of tor?

Thank you for your insightful comments. We understand the confusion with the data presented in Table S1 regarding the copy number estimation of the torA gene in A. flavus. We apologize for not providing a clear explanation for the data in the table. Quantitative real-time PCR (qPCR) is widely used to determine the copy number of a specific gene. It involves amplifying the gene of interest and a reference gene simultaneously using specific primers and probes. By comparing the amplification curves of the gene of interest and the reference gene, we can estimate the relative copy number of the gene.

To address your concern and provide more accurate information, we have re-performed the copy number analysis using southern blot. Southern blot analysis allows for the direct estimation of gene copy number by hybridizing genomic DNA with a specific probe for the gene. This method provides more reliable and accurate results in determining gene copy numbers. We discovered that the A. flavus genome contains a single copy of the torA gene. Consequently, we conducted additional experiments to elucidate its function. Specifically, we generated strains with a xylose-inducible promoter system to modulate the expression of torA (Lines 204-238, page 10).

Point 2: In response: For the knockout of the FRB domain, we used the homologous recombination method, but because tor genes are double-copy genes, there are also double copies in the FRB domain. Despite our efforts, we encountered challenges in precisely determining the location of the other copy of the tor gene. I could not understand these consistent data, why not for using sequencing?

Thank you for your valuable feedback. We determined again and confirmed that the torA gene is a single copy. So we removed this part of the results to avoid any ambiguity or potential misinterpretation.

Point 3: Response in Due to the large number of genes involved, we did not perform a complementation experiment. If there were no complementation data, how to demonstrate data are solid?

Thank you for your important suggestion. We understand that complementation experiments are commonly used to validate gene deletions. Therefore, to ensure the reliability of our data, we have conducted supplementary experiments on specific gene deletions, such as ΔsitA-C and Δppg1-C. Thank you again for your positive comments and valuable suggestions, which have significantly contributed to enhancing the quality of our manuscript (Lines 320-322, page 15).

Point 4: Acknowledge the confusion? We acknowledge the confusion in our presentation and will ensure that accurate genetic nomenclature is used consistently

Thank you for your comments on our manuscript. We recognize the importance of precise and consistent use of genetic nomenclature, as it is critical for the clarity and integrity of our research findings. We have carefully reviewed the sections of our manuscript where genetic terms were used and have made the necessary corrections to ensure that all nomenclature is accurate and used consistently throughout the text.

Point 5: In the revised version of new manuscript, southern blotting was carried out and found only one copy was existed for tested genes at last. Thus, whole manuscript conclusions should be changed. In addition, Reviewer 1 suggestion for using Illumina-sequence strategy, their tor and tapA mutants could be verified whether they are aneuploid?

We would like to express our gratitude for your insightful comments and suggestions. Following the new experimental data obtained from Southern blotting, we have identified that only one copy of the tested genes exists, and we have revised our conclusions throughout the manuscript. This has led to a significant reinterpretation of our results and a reassessment of the implications for our study. Based on this result, we designed and constructed strains with the tor gene under the control of a xylose-inducible promoter. This approach allows for the conditional expression of the tor gene. Thank you once again for your meticulous review (Lines 204-238, page 10).

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The target of rapamycin (TOR) signaling pathway is highly conserved and plays a crucial role in diverse biological processes in eukaryotes. However, the underlying mechanism of the TOR pathway in A. flavus remains elusive. In this study, we identified and characterized seven genes encoding various components of the TOR pathway in A. flavus, and investigated their biological function. The FK506-binding protein Fkbp3 and its lysine succinylation are important for aflatoxin production and rapamycin resistance. As a significant downstream effector molecule of the TorA kinase, the Sch9 kinase regulates aflatoxin B1 (AFB1) synthesis, osmotic and calcium stress response in A. flavus, and this regulation is mediated through its S_TKc, S_TK_X domains, and the ATP binding site at K340. We also showed that the Sch9 kinase may have a regulatory impact on the high-osmolarity glycerol (HOG) signaling pathway. TipA, the other downstream component of the TorA kinase, plays a significant role in regulating sclerotia formation and cell wall stress response in A. flavus. The members of the TapA-phosphatase complexes, SitA and Ppg1, are crucial for various biological processes in A. flavus, including vegetative growth, sclerotia formation, AFB1 biosynthesis, and pathogenicity. Furthermore, we showed that SitA and Ppg1 are involved in regulating lipid droplets (LDs) biogenesis and cell wall integrity (CWI) signaling pathways. In addition, another phosphatase complex, Nem1/Spo7, plays critical roles in hyphal development, conidiation, aflatoxin, and lipid droplets biogenesis. This study provides an important insight into the regulatory network of the TOR signaling pathway and the molecular mechanism of aflatoxin biosynthesis in A. flavus.

Revised version: The target of rapamycin (TOR) signaling pathway is highly conserved and plays a crucial role in diverse biological processes in eukaryotes. Despite its significance, the underlying mechanism of the TOR pathway in A. flavus remains elusive. In this study, we comprehensively analyzed the TOR signaling pathway in A. flavus by identifying and characterizing nine genes that encode distinct components of this pathway. The FK506-binding protein Fkbp3 and its lysine succinylation are important for aflatoxin production and rapamycin resistance. The TorA kinase plays a pivotal role in the regulation of growth, spore production, aflatoxin biosynthesis, and responses to rapamycin and cell membrane stress. As a significant downstream effector molecule of the TorA kinase, the Sch9 kinase regulates aflatoxin B1 (AFB1) synthesis, osmotic and calcium stress response in A. flavus, and this regulation is mediated through its S_TKc, S_TK_X domains, and the ATP binding site at K340. We also showed that the Sch9 kinase may have a regulatory impact on the high-osmolarity glycerol (HOG) signaling pathway. TapA and TipA, the other downstream components of the TorA kinase, play a significant role in regulating cell wall stress response in A. flavus. Moreover, the members of the TapA-phosphatase complexes, SitA and Ppg1, are important for various biological processes in A. flavus, including vegetative growth, sclerotia formation, AFB1 biosynthesis, and pathogenicity. We also demonstrated that SitA and Ppg1 are involved in regulating lipid droplets (LDs) biogenesis and cell wall integrity (CWI) signaling pathways. In addition, another phosphatase complex, Nem1/Spo7, plays critical roles in hyphal development, conidiation, aflatoxin production, and lipid droplets biogenesis. Collectively, our study has provided important insight into the regulatory network of the TOR signaling pathway and has elucidated the underlying molecular mechanisms of aflatoxin biosynthesis in A. flavus.

Introduction

Earlier version: Aspergillus flavus is one of the most important phytopathogenic fungi that frequently contaminate various plants, causing economic losses in agricultural commodities and posing serious health concerns in developing countries [1]. A. flavus is also an opportunistic human pathogenic fungus that can cause aspergillosis in immune-compromised patients [2]. Aflatoxins (AFs), the toxic and carcinogenic secondary metabolites produced by A. flavus, are considered the most potent carcinogen present in nature, causing liver cancer in both humans and animals [3]. The aflatoxin biosynthesis is a complex process involving a series of reactions. Although the majority of the enzyme reactions and genes involved in aflatoxin biosynthesis have been elucidated [4], the related signaling pathway network and specific regulatory mechanisms governing aflatoxin biosynthesis remain elusive. Therefore, exploring the regulatory mechanism of the aflatoxin biosynthesis signaling pathway would provide new insights for preventing A. flavus and aflatoxin contamination.

Revised version: Aspergillus flavus is a highly significant phytopathogenic fungus that frequently contaminates a diverse array of agricultural crops, leading to substantial economic losses in the food and agriculture industry and posing considerable health risks, particularly in developing countries [1]. Moreover, A. flavus is recognized as an opportunistic human pathogenic fungus that can cause aspergillosis in immune-compromised patients [2]. The fungus is notorious for its production of aflatoxins (AFs), which are potent secondary metabolites with toxic and carcinogenic properties and are classified among the most powerful carcinogens known in nature, with established links to liver cancer in both humans and animals [3]. Aflatoxin biosynthesis is a multifaceted and intricate biochemical process encompassing a cascade of enzymatic reactions. Although the majority of the enzyme reactions and genes involved in aflatoxin biosynthesis have been elucidated [4], the related signaling pathway network and specific regulatory mechanisms governing aflatoxin biosynthesis remain elusive. Therefore, exploring the regulatory mechanism of the aflatoxin biosynthesis signaling pathway would provide new insights for preventing A. flavus and aflatoxin contamination.

In eukaryotes, the target of rapamycin (TOR) signaling pathway is highly conserved and plays essential roles in various significant biological processes, including ribosome biosynthesis, cell growth, and autophagy [5]. The serine/threonine-protein kinase Tor serves as a crucial protein component in the TOR signaling pathway, which interacts with other proteins to form multi-protein complexes [6]. In Saccharomyces cerevisiae, TOR exists in two distinct multiprotein functional complexes, namely TORC1 and TORC2 [7]. TORC1-mediated signaling governs cellular growth by modulating various growth-related mechanisms and exhibits sensitivity to rapamycin. Conversely, TORC2-mediated signaling primarily regulates cytoskeletal remodeling but remains unaffected by rapamycin [7]. In yeast, there are two tor genes, tor1 and tor2, whereas in higher eukaryotes such as plants, animals, and filamentous fungi, there is only one tor gene [8–10].

Earlier version: In the fungal kingdom, the TOR signaling pathway has been extensively investigated in three prominent species: the budding yeast S. cerevisiae [11–12], the fission yeast Schizosaccharomyces pombe [13–15], and the human pathogen Candida albicans [16–17]. In S. cerevisiae, macrolide antibiotic rapamycin interacts with the cytosolic protein Fkbp12 to form a complex, and the Fkbp12-rapamycin complex effectively inhibits the activity of the Tor kinase function by binding to the FRB domain [18]. The Tap42 phosphatase complex serves as the primary target of the Tor kinase. Additionally, the Tap42 protein interacts with the Tip41 protein to collaboratively regulate phosphatases, including PP2A and Sit4 [19–20]. Phosphatase Sit4 can dephosphorylate the transcription factor Gln3, regulating the nitrogen source metabolism pathway [21]. In addition, another immediate effector of TORC1 is the AGC kinase Sch9, which participates in various cellular functions processes, such as ribosome biosynthesis, translation initiation, and entry into the G0 phase [22]. The TOR signaling pathway has been documented in various filamentous fungi, including Fusarium graminearum [23–24], Magnaporthe oryzae [25], Phanerochaete chrysosporium [26], Podospora anserine [27], A. nidulans [28], A. fumigatus [29–30], F. fujikuroi [31], Botrytis cinerea [32], F. oxysporum [33], and Mucor circinelloides [34]. However, the TOR signaling pathway has not yet been reported in A. flavus.

Revised version: In the fungal kingdom, the TOR signaling pathway has been extensively investigated in three prominent species: the budding yeast S. cerevisiae [11–12], the fission yeast Schizosaccharomyces pombe [13–15], and the human pathogen Candida albicans [16–17]. In S. cerevisiae, the macrolide antibiotic rapamycin forms a complex with the cytosolic protein Fkbp12, which effectively inhibits the catalytic function of the Tor kinase by interacting with its FRB domain [18]. The Tap42 phosphatase complex serves as the primary target of the Tor kinase and plays a crucial role in the regulation of downstream effectors. Additionally, the Tap42 protein interacts with the Tip41 protein to collaboratively modulate the activities of phosphatases, including PP2A and Sit4 [19–20]. Among these, the Sit4 phosphatase is known to dephosphorylate the transcription factor Gln3, thereby regulating the nitrogen source metabolism pathway [21]. In addition, the AGC kinase Sch9, an immediate effector of TORC1, is involved in regulating various cellular functions processes, such as ribosome biosynthesis, translation initiation, and entry into the G0 phase [22]. The TOR signaling pathway has been documented in various filamentous fungi, including Fusarium graminearum [23–24], Magnaporthe oryzae [25], Phanerochaete chrysosporium [26], Podospora anserine [27], A. nidulans [28], A. fumigatus [29–30], F. fujikuroi [31], Botrytis cinerea [32], F. oxysporum [33], and Mucor circinelloides [34]. However, the TOR signaling pathway in A. flavus remains unexplored.

Earlier version: In eukaryotic cells, the TOR signaling pathway and the mitogen-activated protein kinase (MAPK) signaling pathways play an important role in regulating adaptive responses to extra- and intracellular conditions [35]. Several studies have demonstrated that the TOR signaling pathway interacts with various other signaling pathways, such as MAPK and cell wall integrity (CWI) pathway. These pathways engage in crosstalk, forming a complex metabolic network that regulates cellular processes [36–40]. However, the underlying mechanisms of multiple crosstalks between the TOR and other signaling pathways in A. flavus remain unclear. Therefore, we intended to identify the genes associated with the TOR signaling pathway and elucidate their roles and contributions in regulating vegetative development and aflatoxin biosynthesis in A. flavus. Our results demonstrated that the TOR signaling pathway is involved in multiple cellular processes in A. flavus. Our findings revealed a complex interplay between the TOR pathway and other signaling pathways (such as HOG and CWI) in A. flavus, which play a crucial role in regulating cellular growth and survival under environmental stress conditions.

Revised version: In eukaryotic cells, the TOR signaling pathway and the mitogen-activated protein kinase (MAPK) signaling pathways play an important role in regulating adaptive responses to extra- and intracellular conditions [35]. Several studies have demonstrated that the TOR signaling pathway interacts with various other signaling pathways, such as MAPK and cell wall integrity (CWI) pathway. These pathways are recognized to participate in crosstalk, forming a complex metabolic network that regulates cellular processes [36–40]. However, the underlying mechanisms of multiple crosstalks between the TOR and other signaling pathways in A. flavus remain largely unknown. Therefore, we intended to identify the genes associated with the TOR signaling pathway and elucidate their roles and contributions in regulating vegetative development and aflatoxin biosynthesis in A. flavus. Our results demonstrated that the TOR signaling pathway is involved in multiple cellular processes in A. flavus. Notably, our findings revealed a complex interplay between the TOR pathway and other signaling pathways (such as HOG and CWI) in A. flavus, which play a crucial role in regulating cellular growth and survival under environmental stress conditions.

Results

Rapamycin exhibits inhibitory effects on the growth, sporulation, sclerotia formation, and aflatoxin production in A. flavus

Earlier version: Rapamycin, a secondary metabolite synthesized by Streptomyces hygroscopicus, exhibits efficacy in suppressing specific filamentous fungi. In F. graminearum, rapamycin has a significant inhibitory impact on growth and asexual reproduction [23]. To investigate the effects of rapamycin on vegetative development, sclerotia formation, and secondary metabolism in A. flavus, we conducted an assay to determine the sensitivity of the wild-type strain to rapamycin. We observed that the wild-type strain exhibited a high sensitivity to rapamycin. We found that 100 ng/mL of rapamycin significantly inhibited mycelial growth and conidia formation (Fig. 1A, 1D, 1E). More importantly, the wild-type strain exhibited a significant reduction in sclerotia formation and aflatoxin synthesis (Fig. 1B, 1C, 1F, 1G). Additionally, rapamycin resulted in a significant decrease in spore germination rate and sparser conidiophores (Fig. S1A, S1B). These results illustrated that the TOR signaling pathway may be involved in regulating the growth, development, and secondary metabolism of A. flavus. Calcofluor white is a chemifluorescent blue dye commonly employed for its nonspecific binding affinity to fungal chitin. Fluorescence microscopy analysis revealed a notable reduction in the blue fluorescence on the cell wall following treatment with rapamycin (Fig. S1C). The intracellular lipid droplets within the hyphae were labeled with BODIPY. We observed that the hyphae treated with rapamycin contained more lipid droplets compared to the control group (Fig. S1D). These findings suggest that the TOR pathway may play important roles in maintaining cell wall integrity and facilitating the biogenesis of lipid droplets.

Revised version: Rapamycin, a secondary metabolite synthesized by Streptomyces hygroscopicus, exhibits efficacy in suppressing specific filamentous fungi. In F. graminearum, rapamycin has a significant inhibitory impact on growth and asexual reproduction [23]. To investigate the effects of rapamycin on vegetative development, sclerotia formation, and secondary metabolism in A. flavus, we conducted an assay to determine the sensitivity of the wild-type (WT) strain to rapamycin. We observed that the WT strain exhibited a high sensitivity to rapamycin. We found that 100 ng/mL of rapamycin significantly inhibited mycelial growth and conidia formation (Fig. 1A, 1D, 1E). More importantly, the WT strain exhibited a significant reduction in sclerotia formation and aflatoxin synthesis (Fig. 1B, 1C, 1F, 1G). Additionally, rapamycin resulted in a significant decrease in spore germination rate and sparser conidiophores (Fig. S1A, S1B). These results illustrated that the TOR signaling pathway may be involved in regulating the growth, development, and secondary metabolism of A. flavus. Calcofluor white is a chemifluorescent blue dye commonly employed for its nonspecific binding affinity to fungal chitin. Fluorescence microscopy analysis revealed a notable reduction in the blue fluorescence on the cell wall following treatment with rapamycin (Fig. S1C). The intracellular lipid droplets within the hyphae were labeled with BODIPY. We observed that the hyphae treated with rapamycin contained more lipid droplets compared to the control group (Fig. S1D). These findings suggest that the TOR pathway may play important roles in maintaining cell wall integrity and facilitating the biogenesis of lipid droplets.

Impacts of rapamycin on the growth, sporulation, sclerotia formation,and aflatoxin production of Aspergillus flavus.

(A) Colony morphology of the WT strain cultured on PDA medium amended with different concentrations of rapamycin at 37°C for 5 days.

(B) Colony phenotype of the WT strain grown on CM medium amended with different concentrations of rapamycin at 37°C for 7 days.

(C) TLC assay of AFB1 production by the WT strain cultured in YES liquid media amended with different concentrations of rapamycin at 29°C for 6 days.

(D) Statistical analysis of the colony diameter by the WT strain treated with different concentrations of rapamycin described in (A).

(E) Conidial quantification of the WT strain treated with different concentrations of rapamycin as mentioned in (A).

(F) Quantitative analysis of sclerotium formation by the WT strain treated with different concentrations of rapamycin described in (B).

(G) Relative quantification of AFB1 production by the WT strain treated with different concentrations of rapamycin as mentioned in (C)

(H) Relative expression levels of conidia-related genes in WT strain treated with 100 ng/mL rapamycin after 3, 6, 9 h.

(I) Relative expression levels of sclerotia-related genes in WT strain treated with 100 ng/mL rapamycin after 3, 6, 9 h.

(J) Relative expression levels of aflatoxin biosynthesis regulatory and structural genes in WT strain treated with 100 ng/mL rapamycin after 3, 6, 9 h.

To investigate the potential regulatory role of the TOR pathway in modulating the sporulation, sclerotia formation, and aflatoxin biosynthesis in A. flavus, we examined the expression levels of conidia-related, sclerotia formation-related, and aflatoxin-related genes after 3, 6, and 9 hours of treatment with rapamycin. To determine the effects of rapamycin on conidiation, qRT-PCR was conducted to assess the transcript levels of two conidia-related genes, abaA and brlA. We also conducted qRT-PCR to evaluate the expression levels of three regulators (nsdC, nsdD, and sclR) involved in sclerotia development, and genes involved in AFB1 biosynthesis, such as the regulatory gene aflS, as well as several structural genes (aflC and aflQ). The strains treated with rapamycin exhibited a significant reduction in the expression levels of genes related to sporulation, sclerotia formation, and aflatoxin synthesis (Fig. 1H, 1I, 1J). Overall, these results indicated that the TOR pathway plays a crucial role in regulating the growth, sporulation, sclerotia formation, and aflatoxin biosynthesis in A. flavus by modulating the expression of these genes.

The impact of Fkbp3 and its lysine succinylation on AFB1 biosynthesis and rapamycin resistance in A. flavus

Earlier version: FK506-binding proteins (Fkbps) belong to a highly conserved immunophilin family, which are involved in various cellular functions, such as protein folding, immunosuppression, signaling transduction, and transcription [41]. Fkbps function as molecular switches by binding to target proteins and inducing conformational changes [41]. In S. cerevisiae, Fkbp12 and rapamycin interact to form a complex that exerts a negative regulatory effect on the activity of the Tor kinase [42]. The fkbp12 gene homologue fprA in A. nidulans shares high levels of homology with that from S. cerevisiae [28]. According to the protein sequence of FprA from A. nidulans, four genes encoding putative proteins with FK506 binding domain, named fkbp1 (AFLA_126880), fkbp2 (AFLA_087480), fkbp3 (AFLA_010910), and fkbp4 (AFLA_128200), have been identified from the A. flavus genome. We generated four fkbp knockout strains using homologous recombination, and all fkbp disruption strains were confirmed by PCR and sequencing analysis (Fig. S2). Phenotypic analyses showed that the mycelial growth and conidia formation of all the fkbp mutants were similar to the wild-type strain on the PDA medium (Fig. 2A), suggesting that Fkbps have minimal influence on the process of vegetative development. We evaluated AFB1 synthesis utilizing thin layer chromatography (TLC), and the findings demonstrated a noteworthy reduction in AFB1 synthesis upon deletion of Fkbp3, in comparison to the wild-type strain (Fig. 2B, 2D). Additionally, the quantity of sclerotia exhibited a significant decrease in all fkbp mutants compared to the wild-type strain (Fig. S3A, S3B). The aforementioned data indicated that Fkbp3 plays a crucial role in sclerotia formation and aflatoxin biosynthesis in A. flavus.

Revised version: FK506-binding proteins (Fkbps) belong to a highly conserved immunophilin family, which are involved in various cellular functions, such as protein folding, immunosuppression, signaling transduction, and transcription [41]. Fkbps function as molecular switches by binding to target proteins and inducing conformational changes [41]. In S. cerevisiae, Fkbp12 and rapamycin interact to form a complex that exerts a negative regulatory effect on the activity of the Tor kinase [42]. The fkbp12 gene homologue fprA in A. nidulans shares high levels of homology with that from S. cerevisiae [28]. According to the protein sequence of FprA from A. nidulans, four genes encoding putative proteins with FK506 binding domain, named fkbp1 (AFLA_126880), fkbp2 (AFLA_087480), fkbp3 (AFLA_010910), and fkbp4 (AFLA_128200), have been identified from the A. flavus genome. We generated four fkbp knockout strains using homologous recombination, and all fkbp disruption strains were confirmed by PCR and sequencing analysis (Fig. S2). Phenotypic analyses showed that the mycelial growth and conidia formation of all the fkbp mutants were similar to the WT strain on the PDA medium (Fig. 2A), suggesting that Fkbps have minimal influence on the process of vegetative development. We evaluated AFB1 synthesis utilizing thin layer chromatography (TLC), and the findings demonstrated a noteworthy reduction in AFB1 synthesis upon deletion of Fkbp3, in comparison to the WT strain (Fig. 2B, 2D). Additionally, the quantity of sclerotia exhibited a significant decrease in all fkbp mutants compared to the WT strain (Fig. S3A, S3B). The aforementioned data indicated that Fkbp3 plays a crucial role in sclerotia formation and aflatoxin biosynthesis in A. flavus.

Disruption of fkbp3 significantly increases the resistance of A. flavus to rapamycin and FK506.

(A) Phenotype of the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) grown on PDA amended with rapamycin and FK506 at 37°C for 5 days.

(B) TLC analysis of AFB1 production by the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) cultured in YES liquid medium at 29°C for 6 days.

(C) The growth inhibition rate of the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) under rapamycin and FK506 stress.

(D) AFB1 quantitative analysis of the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) as described in (B).

(E) Phenotype of the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

(F) The growth inhibition rate of the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) under rapamycin stress.

(G) TLC assay of AFB1 production by the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) cultured in YES liquid medium at 29°C for 6 days.

(H) Relative quantification of AFB1 production in the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) as mentioned in (G).

Earlier version: To determine the specific Fkbp responsible for rapamycin sensitivity, all mutant strains were cultured on PDA solid plates amended with 100 ng/mL rapamycin. The strain with disrupted fkbp3 showed significant resistance to rapamycin, while the deletion mutants of fkbp1, fkbp2, and fkbp4 did not exhibit any alteration in resistance to rapamycin (Fig. 2A, 2C). These results demonstrated that Fkbp3 serves as the principal target for rapamycin in A. flavus. In the FK506-sensitivity analysis, we observed a significant suppression of mycelial growth and conidia formation in wild-type strain with the addition of 10 ng/mL FK506 to the PDA medium. Additionally, the deletion of fkbp3 led to the heightened resistance to FK506 compared to the wild-type strain (Fig. 2A, 2C). These findings suggested that Fkbp3 is mainly involved in the modulation of FK506 resistance in A. flavus.

Revised version: To determine the specific Fkbp responsible for rapamycin sensitivity, all mutant strains were cultured on PDA solid plates amended with 100 ng/mL rapamycin. The strain with disrupted fkbp3 showed significant resistance to rapamycin, while the deletion mutants of fkbp1, fkbp2, and fkbp4 did not exhibit any alteration in resistance to rapamycin (Fig. 2A, 2C). These results demonstrated that Fkbp3 serves as the principal target for rapamycin in A. flavus. In the FK506-sensitivity analysis, we observed a significant suppression of mycelial growth and conidia formation in WT strain with the addition of 10 ng/mL FK506 to the PDA medium. Additionally, the deletion of fkbp3 led to the heightened resistance to FK506 compared to the WT strain (Fig. 2A, 2C). These findings suggested that Fkbp3 is mainly involved in the modulation of FK506 resistance in A. flavus.

Earlier version: Previous studies have shown that lysine succinylation plays a significant role in aflatoxin production and pathogenicity in A. flavus [43]. Based on the analysis of our succinylome data, we have identified six reliable succinylation sites (K5, K19, K40, K42, K55, and K65) on the Fkbp3 protein. To identify the function of Fkbp3 succinylation sites, site-directed mutations were constructed according to the homologous recombination strategy. Compared to the wild-type strain, all point mutants exhibited a similar phenotype in vegetative growth and conidiation. Interestingly, we observed that the K19A mutation resulted in increased resistance to rapamycin (Fig. 2E, 2F). Additionally, the K19A strain exhibited a significant reduction in the production of AFB1 compared to the wild-type strain (Fig. 2G, 2H). These findings suggested that the succinylation of Fkbp3 at K19 plays a crucial role in rapamycin resistance and aflatoxin biosynthesis. Furthermore, we found that various point mutants, including K40A, K42A, and K55A, exhibited reduced levels of AFB1 production in comparison to the wild-type strain (Fig. 2G, 2H), suggesting that the succinylation of Fkbp3 at other sites also contributes to AFB1 biosynthesis in A. flavus. The above results indicated that the K19 site likely plays a more prominent role in Fkbp3 succinylation compared to other sites.

Revised version: Previous studies have shown that lysine succinylation plays a significant role in aflatoxin production and pathogenicity in A. flavus [43]. Based on the analysis of our succinylome data, we have identified six reliable succinylation sites (K5, K19, K40, K42, K55, and K65) on the Fkbp3 protein. To identify the function of Fkbp3 succinylation sites, site-directed mutations were constructed according to the homologous recombination strategy. Compared to the WT strain, all point mutants exhibited a similar phenotype in vegetative growth and conidiation. Interestingly, we observed that the K19A mutation resulted in increased resistance to rapamycin (Fig. 2E, 2F). Additionally, the K19A strain exhibited a significant reduction in the production of AFB1 compared to the WT strain (Fig. 2G, 2H). These findings suggested that the succinylation of Fkbp3 at K19 plays a crucial role in rapamycin resistance and aflatoxin biosynthesis. Furthermore, we found that various point mutants, including K40A, K42A, and K55A, exhibited reduced levels of AFB1 production in comparison to the WT strain (Fig. 2G, 2H), suggesting that the succinylation of Fkbp3 at other sites also contributes to AFB1 biosynthesis in A. flavus. The above results indicated that the K19 site likely plays a more prominent role in Fkbp3 succinylation compared to other sites.

Added in the revised version: The TorA kinase plays a critical role in A. flavus

Added in the revised version: In S. cerevisiae, the Tor kinase is known to interact with a diverse range of proteins, resulting in the formation of complex structures known as TORC1 and TORC2 [44]. TORC1 can detect various signals, including nutrients, growth factors, and environmental stress. It plays a crucial role in regulating cellular processes such as gene transcription, protein translation, ribosome synthesis, and autophagy [45]. To identify the ortholog of the Tor protein in A. flavus, the S. cerevisiae Tor protein was utilized as search reference in the A. flavus genome database using the Basic Local Alignment Search Tool (https://blast.ncbi.nlm.nih.gov/Blast.cgi). The genome of A. flavus contains a solitary ortholog of the tor gene (AFLA_044350), which encodes a protein exhibiting 48.91% similarity to S. cerevisiae Tor2. To elucidate the role of torA in A. flavus, we attempted to obtain torA deletion strains by homologous recombination, but not successful. As an alternative approach, a mutant strain with a xylose promoter mutation, referred as xylPtorA, was constructed to validate the function of torA. Subsequent confirmation via PCR analysis ensured the integrity of the mutant strain. Xylan and xylose are highly effective inducers, whereas glucose strongly represses the activity of the promoter. To ascertain the impact of torA on A. flavus, we conducted a comparative analysis of torA transcript levels between the xylPtorA strain and the WT strain. Our findings revealed that the torA transcript levels in the xylPtorA strain were markedly reduced compared to the WT strain (Figure S4A). Both the WT strain and xylPtorA mutant were cultured in YXT medium supplemented with xylose and YGT medium devoid of xylose for a duration of 5 days at 37℃ under dark conditions. Remarkably, the xylPtorA strain exhibited complete inability to thrive in the absence of xylose, with only partial restoration of growth observed upon xylose supplementation (Fig. 3A, 3D). Notably, the xylPtorA strain also demonstrated a markedly reduced growth rate on YXT compared to the WT strain (Fig. 3A, 3C). Thin-layer chromatography (TLC) assays indicated a noteworthy decrease in AFB1 synthesis in the xylPtorA strain compared to the WT strain (Figure 3B, 3F). Additionally, the xylPtorA strain produced fewer conidia and sclerotium formation was absent on YXT medium (Fig. 3E, S5). Interestingly, the xylPtorA strain exhibited a notable increase in sensitivity to rapamycin and SDS stress (Figure 3G, 3H). Collectively, these findings suggested that the TorA kinase plays a crucial role in various biological processes in A. flavus, including vegetative growth, asexual development, sclerotia formation, response to environmental stress, and biosynthesis of aflatoxins.

Added in the revised version: The TorA kinase plays a critical role in A. flavus

Added in the revised version: (A) Phenotype of the WT and the xylPtorA strains in YGT and YXT medium amended with 100 ng/mL rapamycin and 300 mg/mL SDS at 37°C for 5 days.

Added in the revised version: (B) TLC assay of AFB1 production from the WT and the xylPtorA strains cultured in s on YXT medium containing 1 g/L MgSO4·7H2O at 29°C for 6 days.

Added in the revised version: (C) Statistical analysis of the colony diameter by the WT and the xylPtorA strains in YXT medium.

Added in the revised version: (D) Statistical analysis of the colony diameter by the WT and the xylPtorA strains in YGT medium.

Added in the revised version: (E) Conidial quantification of theWT and the xylPtorA strains in YXT medium.

Added in the revised version: (F) Quantitative analysis of AFB1 as shown in (B).

Added in the revised version: (G) The growth inhibition rate of the WT and the xylPtorA strains under rapamycin.

Added in the revised version: (H) The growth inhibition rate of the WT and the xylPtorA strains under SDS stress.

The Sch9 kinase is involved in aflatoxin biosynthesis and the HOG pathway

Earlier version: Sch9, a Ser/Thr kinase belonging to the AGC family, is directly phosphorylated by TORC1. The phosphorylation of Sch9 is inhibited by rapamycin, as well as by carbon or nitrogen starvation [22]. The putative sch9 gene (AFLA_127440) in A. flavus encodes a protein consisting of 706 amino acids, exhibiting 80.31% identity to the Sch9 protein in A. nidulans. To investigate the role of Sch9 in A. flavus, we attempted to create a sch9 deletion mutant using a homologous recombination strategy. The Δsch9 strain exhibited typical vegetative growth and conidiation similar to the wild-type strain (Fig. 3B). TLC assay and quantitative analysis showed a significantly decreased aflatoxin production in the Δsch9 strain compared to the wild-type strain (Fig. 3C, 3G). These results indicated that Sch9 regulates aflatoxin biosynthesis in A. flavus.

Revised version: Sch9, a Ser/Thr kinase belonging to the AGC family, is directly phosphorylated by TORC1. The phosphorylation of Sch9 is inhibited by rapamycin, as well as by carbon or nitrogen starvation [22]. The putative sch9 gene (AFLA_127440) in A. flavus encodes a protein consisting of 706 amino acids, exhibiting 80.31% identity to the Sch9 protein in A. nidulans. To investigate the role of Sch9 in A. flavus, we attempted to create a sch9 deletion mutant using a homologous recombination strategy. The Δsch9 strain exhibited typical vegetative growth and conidiation similar to the WT strain (Fig. 4B). TLC assay and quantitative analysis showed a significantly decreased aflatoxin production in the Δsch9 strain compared to the WT strain (Fig. 4C, 4G). These results indicated that Sch9 regulates aflatoxin biosynthesis in A. flavus.

The Sch9 kinase participates in aflatoxin biosynthesis and the HOG pathway.

(A) The structure diagram of the Sch9 kinase.

(B) Phenotype of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc, sch9ΔS_TK_X, and sch9K340A strains grown on PDA medium amended with 100 ng/mL rapamycin at 37°C for 5 days.

(C) TLC assay of AFB1 production from the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc, sch9ΔS_TK_X, and sch9K340A strains cultured in YES liquid medium at 29°C for 6 days.

(D) The phenotypes of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc, sch9ΔS_TK_X, and sch9K340Astrains on YGT media amended with 1 M NaCl and 1 M KCl for 3 days.

(E) The phosphorylation levels of Hog1 in the WT and Δsch9 strains were determined with or without osmotic stress.

(F) The growth inhibition rate of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc, sch9ΔS_TK_X, and sch9K340A strains under rapamycin stress.

(G) Quantitative analysis of AFB1 as shown in (C).

(H) The growth inhibition rate of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc, sch9ΔS_TK_X, and sch9K340A strains under osmotic stress.

Earlier version: In yeast, Sch9 plays an important role in the transcriptional activation of osmostress inducible genes [44]. To characterize the roles of Sch9 in stress response, we examined the sensitivity of the sch9 deletion strain to various stresses including osmotic stress and calcium stress. Compared with the wild-type strain, the Δsch9 mutant significantly increased sensitivity to calcium stress (Fig. S4A, S4B), and decreased sensitivity to osmotic stress induced by NaCl and KCl (Fig. 3D, 3H). In A. flavus, the Sch9 kinase consists of several domains, including the protein kinase C conserved region 2 (C2), protein kinase (S_TKc), AGC-kinase C-terminal (S_TK_X) domain, and an ATP binding site (Fig. 3A). To explore the role of these domains and the ATP binding site in Sch9 kinase, we constructed deletion strains for both domains and a sch9K340A mutant. The results revealed that the sch9ΔS_TKc strain, sch9ΔS_TK_X strain, and sch9K340A mutant exhibited similar phenotypes in osmotic stress and AFB1 production as the Δsch9 mutant (Fig. 3B, 3C, 3D). Therefore, we speculate that the protein kinase domain, kinase C-terminal domain, and the ATP binding site at K340 play a crucial role in modulating the impact of Sch9 on aflatoxin biosynthesis and stress response in A. flavus.

Revised version: In yeast, Sch9 plays an important role in the transcriptional activation of osmostress inducible genes [46]. To characterize the roles of Sch9 in stress response, we examined the sensitivity of the sch9 deletion strain to various stresses including osmotic stress and calcium stress. Compared with the WT strain, the Δsch9 mutant significantly increased sensitivity to calcium stress (Fig. S6A, S6B), and decreased sensitivity to osmotic stress induced by NaCl and KCl (Fig. 4D, 4H). In A. flavus, the Sch9 kinase consists of several domains, including the protein kinase C conserved region 2 (C2), protein kinase (S_TKc), AGC-kinase C-terminal (S_TK_X) domain, and an ATP binding site (Fig. 4A). To explore the role of these domains and the ATP binding site in Sch9 kinase, we constructed deletion strains for both domains and a sch9K340A mutant. The results revealed that the sch9ΔS_TKc strain, sch9ΔS_TK_X strain, and sch9K340A mutant exhibited similar phenotypes in osmotic stress and AFB1 production as the Δsch9 mutant (Fig. 4B, 4C, 4D). Therefore, we speculate that the protein kinase domain, kinase C-terminal domain, and the ATP binding site at K340 play a crucial role in modulating the impact of Sch9 on aflatoxin biosynthesis and stress response in A. flavus.

Earlier version: To investigate the potential function of Sch9 in the HOG and TOR pathways in A. flavus, we conducted a sensitivity analysis of the Δsch9 strain to rapamycin. Compared to the wild-type strain, the sch9 deletion strain exhibited a slightly higher sensitivity to rapamycin (Fig. 3B, 3F). The MAKP Hog1 is an element of the high osmolarity glycerol (HOG) pathway. To verify the connection between Sch9 and the HOG-MAPK pathway, the phosphorylation levels of Hog1 kinase were measured under osmotic stress (1M NaCl). Western blot analysis showed that the phosphorylation level of Hog1 increased significantly in the Δsch9 strain (Fig. 3E). This finding suggested that Sch9 plays a pivotal role in modulating the response of Hog1 to osmotic stress. Taken together, these results reflected that Sch9 regulates osmotic stress response via the HOG pathway in A. flavus.

Revised version: To investigate the potential function of Sch9 in the HOG and TOR pathways in A. flavus, we conducted a sensitivity analysis of the Δsch9 strain to rapamycin. Compared to the WT strain, the sch9 deletion strain exhibited a slightly higher sensitivity to rapamycin (Fig. 4B, 4F). The MAKP Hog1 is an element of the high osmolarity glycerol (HOG) pathway. To verify the connection between Sch9 and the HOG-MAPK pathway, the phosphorylation levels of Hog1 kinase were measured under osmotic stress (1M NaCl). Western blot analysis showed that the phosphorylation level of Hog1 increased significantly in the Δsch9 strain (Fig. 4E). This finding suggested that Sch9 plays a pivotal role in modulating the response of Hog1 to osmotic stress. Taken together, these results reflected that Sch9 regulates osmotic stress response via the HOG pathway in A. flavus.

Earlier version: TipA participate in regulation of sclerotia development and cell wall stress response

Revised version: TapA and TipA participate in the regulation of cell wall stress response

Added in the revised version: Tap42, a protein that associates with phosphatase 2A, serves as the direct target of the Tor kinase in yeast. Tap42 plays a crucial role in various Tor functions, particularly in the regulation of transcriptional processes [47]. The genome of A. flavus contains a homolog of tap42 named tapA (AFLA_092770), which is predicted to encode a 355 amino-acid protein (sharing 31.9% identity with S. cerevisiae Tap42). To further determine the function of the TapA regulator in the TOR signaling pathway of A. flavus, we used homologous recombination to replace the original promoter with the gpdA promoter, a well-characterized strong promoter, to construct an overexpression mutant (OE::tapA). The transformants of the mutant were verified using PCR. Quantitative expression analysis of the tapA gene in the overexpressed strains revealed a substantial upregulation compared to the WT strain (Fig. S4B). Our results showed that the colony diameter of the OE::tapA mutant was not significantly changed compared to the WT strain (Fig. 5A). However, the rapamycin sensitivity assay indicated that the OE::tapA strain displayed increased sensitivity to rapamycin (Fig. 5A, 5D). We found no significant difference in aflatoxin production between the OE::tapA and the WT strains (Fig. 5C, 5E), suggesting that TapA may not play a direct role in the regulation of aflatoxin biosynthesis. In addition, the OE::tapA strain was more sensitive to cell wall and cell membrane stress, including SDS, CFW, and CR (Fig. 5B, 5F). These findings have demonstrated the significantly positive impact of TapA on cell wall stress in A. flavus.

Removed in the revised version: In S. cerevisiae, Tip41 was identified as a Tap42-interacting protein and has been measured to function as a negative regulator of Tap42, which downregulates TORC1 signaling via activation of PP2A phosphatase [45]. The putative orthologue of tip41 in A. flavus is named tipA (AFLA_047310). The amino acid sequence of TipA exhibits a similarity of 40.49% to Tip41 in S. cerevisiae. To elucidate the function of the TOR signaling pathway protein TipA in A. flavus, we generated the tipA deletion mutants. The ΔtipA strain exhibited a slightly reduced number of conidia produced on PDA compared to the wild-type strain (Fig. 4A, 4C). Meanwhile, the ΔtipA strain exhibited decreased sensitivity to rapamycin (Fig. 4A, 4E). The ΔtipA strain was unable to produce any sclerotia (Fig. 4B, 4F). Whereas ΔtipA strain do not affect aflatoxin production(Fig. 4D, 4G). In addition, the ΔtipA strain was more sensitive to the cell wall and cell membrane stress, including CR and SDS (Fig. 4H, 4I). Additionally, the expression levels of nsdC and nsdD exhibited a decrease in ΔtipA strain compared to the wild-type strain (Fig. 4J). These results indicated that TipA is crucial for the sclerotial formation and cell wall stress in A. flavus.

Earlier version: TipA regulates clerotia development and cell wall stress in A. flavus.

Revised version: TapA regulates cell wall stress in A. flavus.

Added in the revised version: (A) Phenotype of the WT and OE::tapA strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Added in the revised version: (B) Colony morphology of the WT and OE::tapA strains grown on PDA media supplemented with 200 µg/mL CR, 200µg/mL CFW, or 300 µg/mL SDS at 37°C for 5 days.

Added in the revised version: (C) TLC assay of AFB1 production by the WT and OE::tapA strains cultured in YES liquid medium at 29°C for 6 days.

Added in the revised version: (D) The growth inhibition rate of the WT and OE::tapA strains under rapamycin.

Added in the revised version: (E) Quantitative analysis of AFB1 as shown in (C).

Added in the revised version: (F) The growth inhibition rate of the WT and OE::tapA strains under cell wall and cell membrane stress.

Added in the revised version: In S. cerevisiae, Tip41 was identified as a Tap42-interacting protein and has been measured to function as a negative regulator of Tap42, which downregulates TORC1 signaling via activation of PP2A phosphatase [48]. The putative orthologue of tip41 in A. flavus is named tipA (AFLA_047310). The amino acid sequence of TipA exhibits a similarity of 40.49% to Tip41 in S. cerevisiae. To elucidate the function of the TOR signaling pathway protein TipA in A. flavus, we generated the tipA deletion mutants. The ΔtipA strain exhibited a slightly reduced number of conidia produced on PDA compared to the WT strain (Fig. 6A, 6C). Meanwhile, the ΔtipA strain exhibited decreased sensitivity to rapamycin (Fig. 6A, 6E). The ΔtipA strain was unable to produce any sclerotia (Fig. 6B, 6F). Whereas ΔtipA strain do not affect aflatoxin production(Fig. 6D, 6G). In addition, the ΔtipA strain was more sensitive to the cell wall and cell membrane stress, including CR and SDS (Fig. 6H, 6I). Additionally, the expression levels of nsdC and nsdD exhibited a decrease in ΔtipA strain compared to the WT strain (Fig. 6J). These results indicated that TipA is crucial for the sclerotial formation and cell wall stress in A. flavus.

Added in the revised version: TipA regulates sclerotia development and cell wall stress in A. flavus.

(A) Phenotype of the WT and ΔtipA strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

(B) Phenotypic characterization of the WT and ΔtipA strains grown on CM medium at 37°C for 7 days.

(C) Morphology of conidiophores in the WT and ΔtipA strains observed by microscope.

(D) TLC assay of AFB1 production by the WT and ΔtipA strains cultured in YES liquid medium at 29°C for 6 days.

(E) The growth inhibition rate of the WT and ΔtipA strains under rapamycin stress.

(F) Amount of sclerotia produced by the WT and ΔtipA strains.

(G) AFB1 quantitative analysis of the WT and ΔtipA strains described in (D).

(H) Colony morphology of the WT and ΔtipA strains grown on PDA media supplemented with 200 µg/mL CR, 200µg/mL CFW, and 100 µg/mLSDS at 37°C for 5 days.

(I) The growth inhibition rate of the WT and ΔtipA strains under cell wall and cell membrane stress.

(J) Relative expression levels of sclerotia formation genes in the WT and ΔtipA strains.

Phosphatase SitA and Ppg1 are involved in growth, conidiation, and sclerotial formation

Earlier version: In S. cerevisiae, Tap42 has been identified to interact with the catalytic subunit of type 2A protein phosphatases, including Pph3, Pph21, and Pph22, as well as the type 2A-like phosphatases Sit4 and Ppg1[19]. Phosphatase sitA (AFLA_108450) and ppg1 (AFLA_132610) were identified in A. flavus, encoding putative orthologs of sit4 and ppg1 in S. cerevisiae. To determine the biological functions of phosphatase SitA and Ppg1 in A. flavus, we disrupted the sitA and ppg1 by homologous recombination strategy. To further investigate the relationship between phosphatase SitA and Ppg1, we additionally generated complemented and double deletion mutants. The vegetative development phenotype analysis showed that the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains grew significantly more slowly than the wild-type strain on PDA (Fig. 5A, 5D). Additionally, the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains were more sensitive to rapamycin (Fig. 5A, 5E). In addition, the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains formed lower number of conidiophores (Fig. 5A ). To gain further insight into the role of SitA and Ppg1 in the process of conidiation, qRT-PCR was employed to determine the transcript levels of conidiation-related genes (abaA and brlA). The expression levels of the abaA and brlA genes were significantly downregulated in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants when compared to the wild-type and complemented strains (Fig. 5G). The analysis of sclerotial formation revealed that the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants exhibited a complete inability to produce any sclerotia (Fig. 5B, 5E). To further prove these findings, we conducted qRT-PCR to determine the expression levels of three regulators (nsdC, nsdD, and sclR) of sclerotia development. Our findings indicated that the expression levels of nsdC, nsdD, and sclR were all significantly reduced in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants (Fig. 5H), which is consistent with the observed decrease in sclerotia production in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. These results showed that SitA and Ppg1 are crucial for hyphal development, conidiation, and sclerotia formation in A. flavus.

Revised version: In S. cerevisiae, Tap42 has been identified to interact with the catalytic subunit of type 2A protein phosphatases, including Pph3, Pph21, and Pph22, as well as the type 2A-like phosphatases Sit4 and Ppg1[19]. Phosphatase sitA (AFLA_108450) and ppg1 (AFLA_132610) were identified in A. flavus, encoding putative orthologs of sit4 and ppg1 in S. cerevisiae. To determine the biological functions of phosphatase SitA and Ppg1 in A. flavus, we disrupted the sitA and ppg1 by homologous recombination strategy. To further investigate the relationship between phosphatase SitA and Ppg1, we additionally generated complemented and double deletion mutants. The vegetative development phenotype analysis showed that the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains grew significantly more slowly than the WT strain on PDA (Fig. 7A, 7D). Additionally, the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains were more sensitive to rapamycin (Fig. 7A, 7E). In addition, the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains formed lower number of conidiophores (Fig. 7A). To gain further insight into the role of SitA and Ppg1 in the process of conidiation, qRT-PCR was employed to determine the transcript levels of conidiation-related genes (abaA and brlA). The expression levels of the abaA and brlA genes were significantly downregulated in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants when compared to the WT and complemented strains (Fig. 7G). The analysis of sclerotial formation revealed that the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants exhibited a complete inability to produce any sclerotia (Fig. 7B, 7E). To further prove these findings, we conducted qRT-PCR to determine the expression levels of three regulators (nsdC, nsdD, and sclR) of sclerotia development. Our findings indicated that the expression levels of nsdC, nsdD, and sclR were all significantly reduced in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants (Fig. 7H), which is consistent with the observed decrease in sclerotia production in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. These results showed that SitA and Ppg1 are crucial for hyphal development, conidiation, and sclerotia formation in A. flavus.

The impact of the phosphatases SitA and Ppg1 on the growth, conidiation, and aflatoxin biosynthesis in A. flavus.

(A) Colony morphology of the WT, single knockout, and double knockout strains cultured on PDA medium amended with 100 ng/mL rapamycin at 37℃ for 5 days.

(B) Colony morphology of the WT, single knockout, and double knockout strains cultured on CM medium at 37℃ for 7 days.

(C) TLC analysis of AFB1 production from the WT, single knockout, and double knockout strains cultured in YES liquid medium at 29℃ for 6 days.

(D) Growth diameter of the WT, single knockout, and double knockout strains on PDA media.

(E) The growth inhibition rate of the WT, single knockout, and double knockout strains under rapamycin stress.

(F) AFB1 quantitative analysis of the WT, single knockout, and double knockout strains.

(G) Relative expression levels of conidia synthesis genes in the WT, single knockout, and double knockout strains.

(H) Relative expression levels of sclerotia synthesis genes in the WT, single knockout, and double knockout strains.

(I) Relative expression levels of aflatoxin biosynthesis genes in the WT, single knockout, and double knockout strains.

SitA and Ppg1 play crucial roles in aflatoxin biosynthesis and pathogenicity

Earlier version: We assayed AFB1 biosynthesis in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. TLC analyses revealed that the ΔsitA and ΔsitA/Δppg1 strains exhibited a decrease in AFB1 production, whereas the Δppg1 strain showed a significant increase in AFB1 production compared to the wild-type and complemented strains (Fig. 5C, 5F). The expression levels of genes related to aflatoxin synthesis, including structural genes (aflO and aflP) and regulatory genes (aflR and aflS), were determined by qRT-PCR in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. The results revealed a significant downregulation of these genes in the ΔsitA and ΔsitA/Δppg1 mutants, while upregulated in the Δppg1 mutants, compared to the wild-type and complemented strains (Fig. 5I). These results indicated that SitA and Ppg1 play different roles in regulating aflatoxin biosynthesis in A. flavus.

Revised version: We assayed AFB1 biosynthesis in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. TLC analyses revealed that the ΔsitA and ΔsitA/Δppg1 strains exhibited a decrease in AFB1 production, whereas the Δppg1 strain showed a significant increase in AFB1 production compared to the WT and complemented strains (Fig. 7C, 7F). The expression levels of genes related to aflatoxin synthesis, including structural genes (aflO and aflP) and regulatory genes (aflR and aflS), were determined by qRT-PCR in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. The results revealed a significant downregulation of these genes in the ΔsitA and ΔsitA/Δppg1 mutants, while upregulated in the Δppg1 mutants, compared to the WT and complemented strains (Fig. 7I). These results indicated that SitA and Ppg1 play different roles in regulating aflatoxin biosynthesis in A. flavus.

Earlier version: We conducted pathogenicity tests on peanuts and maize to examine the pathogenicity of the phosphatases SitA and Ppg1 on crops. The results showed that the wild-type and complemented strains demonstrated complete virulence on all peanut and maize seeds. In contrast, the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants displayed compromised colonization on peanut and maize seeds (Fig. 6A). In addition, the number of conidial productions in these mutants on the infected peanut seeds also dramatically decreased compared to those of the wild-type and complemented strains (Fig. 6C). The thin-layer chromatography analyses revealed that the Δppg1 mutant exhibited increased production of AFB1 in peanuts and maize. Conversely, the levels of aflatoxin produced by the ΔsitA and ΔsitA/Δppg1 mutants on peanut and maize were significantly reduced compared to the wild-type and complemented strains (Fig. 6B, 6D). This observation aligns with the above results of aflatoxin biosynthesis from all mutant strains in the YES medium. All these data showed that SitA and Ppg1 are crucial for crop seeds’ pathogenicity.

Revised version: We conducted pathogenicity tests on peanuts and maize to examine the pathogenicity of the phosphatases SitA and Ppg1 on crops. The results showed that the WT and complemented strains demonstrated complete virulence on all peanut and maize seeds. In contrast, the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants displayed compromised colonization on peanut and maize seeds (Fig. 8A). In addition, the number of conidial productions in these mutants on the infected peanut seeds also dramatically decreased compared to those of the WT and complemented strains (Fig. 8C). The thin-layer chromatography analyses revealed that the Δppg1 mutant exhibited increased production of AFB1 in peanuts and maize. Conversely, the levels of aflatoxin produced by the ΔsitA and ΔsitA/Δppg1 mutants on peanut and maize were significantly reduced compared to the WT and complemented strains (Fig. 8B, 8D). This observation aligns with the above results of aflatoxin biosynthesis from all mutant strains in the YES medium. All these data showed that SitA and Ppg1 are crucial for crop seeds’ pathogenicity.

Pathogenicity analysis of the phosphatase SitA and Ppg1 in A. flavus.

(A) Phenotypes of peanut and maize infected by the WT, single knockout, and double knockout strains.

(B) TLC was used to detect the AFB1 from peanut and maize infected by the WT, single knockout, and double knockout strains.

(C) Quantitative analysis of conidia production from peanut and maize infected by the WT, single knockout, and double knockout strains.

(D) Quantitative analysis of AFB1 from peanut and maize infected by the WT, single knockout, and double knockout strains.

SitA and Ppg1 are involved in regulating cell wall integrity

Earlier version: To explore the role of phosphatases SitA and Ppg1 in response to cell wall stress, all strains were cultivated on PDA amended with CFW, CR, and SDS. We found that the relative growth inhibition of ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants induced by cell wall stress was significantly increased compared to the wild-type and complemented strains (Fig. 7A, 7B). Further investigations using qRT-PCR revealed a notable reduction in the expression levels of genes related to the synthesis of cell walls, specifically chitin synthase genes chsA, chsD, and chsE in the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains (Fig. 7C). Western blot analysis also revealed that the phosphorylation level of Slt2 was significantly elevated in the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains under cell wall stress (200 μg/ml CFW) (Fig. 7D). In conclusion, the phosphatases SitA and Ppg1 play a crucial role in regulating cell wall integrity.

Revised version: To explore the role of phosphatases SitA and Ppg1 in response to cell wall stress, all strains were cultivated on PDA amended with CFW, CR, and SDS. We found that the relative growth inhibition of ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants induced by cell wall stress was significantly increased compared to the WT and complemented strains (Fig. 9A, 9B). Further investigations using qRT-PCR revealed a notable reduction in the expression levels of genes related to the synthesis of cell walls, specifically chitin synthase genes chsA, chsD, and chsE in the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains (Fig. 9C). Western blot analysis also revealed that the phosphorylation level of Slt2 was significantly elevated in the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains under cell wall stress (200 μg/ml CFW) (Fig. 9D). In conclusion, the phosphatases SitA and Ppg1 play a crucial role in regulating cell wall integrity.

Sensitivity of phosphatase SitA and Ppg1 to cell wall damaging agents in A. flavus.

(A) Morphology of the WT, single knockout, and double knockout strains grown on YES media supplemented with 200 µg/mL CR, 200 µg/mL CFW, or 100 µg/mL SDS at 37°C for 5 days.

(B) The growth inhibition rate of the WT, single knockout, and double knockout strains under cell wall and cell membrane stress.

(C) Relative expression levels of cell wall synthesis genes (chsA, chsD and chsE) in the WT, single knockout, and double knockout strains.

(D) The phosphorylation level of Slt2 in the WT, single knockout, and double knockout strains was detected with or without cell wall stress.

The phosphatase complex Nem1/Spo7 is important for growth, aflatoxin biosynthesis, and LD biogenesis

Earlier version: Lipid droplets (LDs) are highly dynamic spherical organelles, which appear in the cytosol of most eukaryotic cell types[46]. Rapamycin treatment led to the accumulation of lipid droplets in various filamentous fungi [24]. The FgTOR pathway regulates the phosphorylation status of FgNem1 mainly via the FgPpg1/Sit4-FgCak1/other kinase cascade in F. graminearum [24]. In this study, we observed that the treatment of hyphae with rapamycin resulted in the accumulation of LD biogenesis (Fig. 8A). This finding suggested that the TOR pathway may regulate the LD biogenesis. However, the underlying regulatory mechanism of LD biosynthesis remain largely unknown in A. flavus. To investigate the regulatory mechanism of the TOR pathway in A. flavus that controls LD biogenesis, we examined LD accumulation in Δsch9, ΔsitA, and Δppg1 mutants. BODIPY staining assays revealed that LD biogenesis was not observed in the ΔsitA and Δppg1 strains upon treatment with rapamycin, in contrast to the wild-type and Δsch9 strains (Fig. 8A, 8B, 8C, 8D). These results indicated that LD biogenesis induced by rapamycin is predominantly reliant on the phosphatases SitA and Ppg1 in A. flavus. In F. graminearum, the phosphatase FgNem1/FgSpo7 complex was involved in the induction of LD biogenesis by rapamycin[24]. Phosphatase nem1 (AFLA_002649) and spo7 (AFLA_028590) were identified in A. flavus, which encode putative orthologs of nem1 and spo7 in F. graminearum. To examine the role of the nem1 and spo7 genes in A. flavus, we generated Δnem1 and Δspo7 mutants using homologous recombination.

Revised version: Lipid droplets (LDs) are highly dynamic spherical organelles, which appear in the cytosol of most eukaryotic cell types[49]. Rapamycin treatment led to the accumulation of lipid droplets in various filamentous fungi [24]. The FgTOR pathway regulates the phosphorylation status of FgNem1 mainly via the FgPpg1/Sit4-FgCak1/other kinase cascade in F. graminearum [24]. In this study, we observed that the treatment of hyphae with rapamycin resulted in the accumulation of LD biogenesis (Fig. 10A). This finding suggested that the TOR pathway may regulate the LD biogenesis. However, the underlying regulatory mechanism of LD biosynthesis remain largely unknown in A. flavus. To investigate the regulatory mechanism of the TOR pathway in A. flavus that controls LD biogenesis, we examined LD accumulation in Δsch9, ΔsitA, and Δppg1 mutants. BODIPY staining assays revealed that LD biogenesis was not observed in the ΔsitA and Δppg1 strains upon treatment with rapamycin, in contrast to the WT and Δsch9 strains (Fig. 10A, 10B, 10C, 10D). These results indicated that LD biogenesis induced by rapamycin is predominantly reliant on the phosphatases SitA and Ppg1 in A. flavus. In F. graminearum, the phosphatase FgNem1/FgSpo7 complex was involved in the induction of LD biogenesis by rapamycin[24]. Phosphatase nem1 (AFLA_002649) and spo7 (AFLA_028590) were identified in A. flavus, which encode putative orthologs of nem1 and spo7 in F. graminearum. To examine the role of the nem1 and spo7 genes in A. flavus, we generated Δnem1 and Δspo7 mutants using homologous recombination. The Δnem1 and Δspo7 mutants displayed a notable reduction in both vegetative growth and conidiation compared to the WT strain (Fig. 11A, 11C, 11D). TLC assay revealed a significantly decreased aflatoxin production in the Δnem1 and Δspo7 compared to the WT strain (Fig. 11B). In addition, BODIPY staining showed no visible lipid droplets in the hyphae of Δnem1 and Δspo7 after rapamycin treatment (Fig. 9E). Taken together, these results indicated that Nem1/Spo7 are involved in vegetative growth, conidiation, aflatoxin production, and LD biogenesis.

Phosphatase SitA and Ppg1 are involved in regulating lipid droplet biogenesis.

(A) The intracellular lipid droplets were stained by boron dipyrromethene difluoride (BODIPY) in hyphae of the WT strain and observed under fluorescence microscopy. The phenotype of lipid droplets accumulation in the mycelia of the WT strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

(B) Phenotype of lipid droplets accumulation in the mycelia of the Δsch9 strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

(C) Phenotype of lipid droplets accumulation in the mycelia of the ΔsitA strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

(D) Phenotype of lipid droplets accumulation in the mycelia of the Δppg1 strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

Removed in the revised version: The Δnem1 and Δspo7 mutants displayed a notable reduction in both vegetative growth and conidiation compared to the wild-type strain (Fig. 9A, 9C, 9D). TLC assay revealed a significantly decreased aflatoxin production in the Δnem1 and Δspo7 compared to the wild-type strain (Fig. 9B). In addition, BODIPY staining showed no visible lipid droplets in the hyphae of Δnem1 and Δspo7 after rapamycin treatment (Fig. 9E). Taken together, these results indicated that Nem1/Spo7 are involved in vegetative growth, conidiation, aflatoxin production, and LD biogenesis.

Phosphatase complex Nem1/Spo7 play significant role in growth, conidiation, aflatoxin, and lipid droplet biogenesis.

(A) Phenotype of the WT, Δnem1, and Δspo7 strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

(B) TLC assay of AFB1 production from the WT, Δnem1, and Δspo7 strains cultured in YES liquid medium at 29°C for 6 days.

(C) Growth diameter of the WT, Δnem1, and Δspo7 strains on PDA media.

(D) Statistical analysis of the sporulation by the WT, Δnem1, and Δspo7 strains.

(E) Phenotype of lipid droplets accumulation in the mycelia of the WT, Δnem1, and Δspo7 strains treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

Removed in the revised version: Discussion

Removed in the revised version: Aflatoxins (AFs) are a class of toxic secondary metabolites synthesized by A. flavus, including aflatoxin B1, B2, G1, and G2, among which AFB1 has been regarded as the most pathogenic and toxic aflatoxin[47]. Therefore, understanding the regulatory mechanisms involved in fungal secondary metabolites biosynthesis would be instrumental for effectively addressing aflatoxin control in A. flavus. In this study, we demonstrated that the target of rapamycin (TOR) signaling pathway plays a pivotal role in regulating growth, sporulation, sclerotia formation, aflatoxin production, and various stress responses in A. flavus. Through conducting functional studies on genes within the TOR signaling pathway, we explored the conservation and complexity exhibited by the TOR signaling pathway in A. flavus. In addition, we constructed a crosstalk network between the TOR and other signaling pathways (HOG, CWI) and analyzed the specific regulatory process. We proposed that the TOR signaling pathway interacts with multiple signaling pathways, including HOG and CWI pathways, to modulate the cellular responses to diverse environmental stresses (Fig. 10).

The proposed model of the target of rapamycin (TOR) pathway in A. flavus.

Based on the above results, we propose a hypothesis regarding the TOR signaling pathway model in A. flavus. Rapamycin forms a complex with Fkbp3, and this particular complex can bind to the TorA kinase, thereby impeding its regular functionality. Sch9, functioning as a downstream element of the TorA kinase, regulates aflatoxin biosynthesis and HOG signaling pathway. As another important target of the TorA kinase, TapA-phosphatase complex is involved in the regulation of growth, sporulation, sclerotia, aflatoxin production, CWI signaling pathway, lipid droplet synthesis, and other processes. In conclusion, the TOR signaling pathway plays a crucial role in various aspects of A. flavus, including vegetative development, stress response, aflatoxin biosynthesis, and pathogenicity.

Added in the revised version: Discussion

Added in the revised version: Aflatoxins (AFs) are a class of toxic secondary metabolites synthesized by A. flavus, including aflatoxin B1, B2, G1, and G2, among which AFB1 has been regarded as the most pathogenic and toxic aflatoxin[50]. Therefore, understanding the regulatory mechanisms involved in fungal secondary metabolites biosynthesis would be instrumental for effectively addressing aflatoxin control in A. flavus. In this study, we demonstrated that the target of rapamycin (TOR) signaling pathway plays a pivotal role in regulating growth, sporulation, sclerotia formation, aflatoxin production, and various stress responses in A. flavus. Through conducting functional studies on genes within the TOR signaling pathway, we explored the conservation and complexity exhibited by the TOR signaling pathway in A. flavus. In addition, we constructed a crosstalk network between the TOR and other signaling pathways (HOG, CWI) and analyzed the specific regulatory process. We proposed that the TOR signaling pathway interacts with multiple signaling pathways, including HOG and CWI pathways, to modulate the cellular responses to diverse environmental stresses (Fig. 10).

Earlier version: We have observed that rapamycin has a significant impact on the growth and sporulation of A. flavus (Fig. 1A). Additionally, it exhibits a strong inhibitory effect on sclerotia production and aflatoxin biosynthesis (Fig. 1B, 1C). The inhibitory effect has been documented in various fungal species, including Candida albicans [48], Cryptococcus neoformans [48], Podospora anserine [27], A. nidulans [28], Magnaporthe oryzae [39], and F. graminearum [23]. This suggested that the TOR signaling pathway may have a conserved function in filamentous fungi. The number of FK506-binding proteins varies among different species: there are four Fkbps proteins in S. cerevisiae and three Fkbps proteins in S. pombe [49]. In S. cerevisiae, rapamycin can bind to Fkbp protein, resulting in the irreversible inhibition of the G1 phase of the cell cycle and the regulation of cell growth [50]. In A. nidulans, the fkbp12 gene homologue fprA displays considerable sequence similarity with homologues identified in the S. cerevisiae [28]. In this study, based on the homologous alignment of the protein sequence of FprA, it was observed that four fkbp genes are present in A. flavus. Among these, deletion of fkbp3 exhibited enhanced tolerance to rapamycin (Fig. 2A). So we speculated that rapamycin forms a complex with Fkbp3 in A. flavus, subsequently targeting downstream target genes to exert its biological function. We found that Fkbp3 is involved in sclerotia formation and aflatoxin biosynthesis in A. flavus (Fig. 2A, S3), and subsequent analysis revealed that the underlying mechanism may be associated with succinylation modification. The succinylation sites of Fkbp3 were identified through the utilization of succinylation proteomics data. By constructing point mutant strains and conducting phenotype experiments, we found that the succinylation of Fkbp3 at residue K19 is involved in rapamycin resistance and aflatoxin biosynthesis (Fig. 2E, 2G). Previous studies have demonstrated that Fkbp12 has been identified as the receptor for FK506 and rapamycin, which inhibits calcineurin and target rapamycin complex 1 (TORC1), respectively [50]. The deletion of Fkbp12 confers the resistance to rapamycin in C. albicans [48] and C. neoformans [48]. The entomopathogenic fungus Beauveria bassiana contains three putative fkbp genes, and disruption of Bbfkbp12 significantly increases the resistance to rapamycin and FK506 [51]. The plant pathogenic fungus F. graminearum harbors three Fkbps proteins, among which deletion of Fgfkbp12 leads to resistance to rapamycin [23]. Thus, we propose that these Fkbps may exhibit analogous functionalities in conferring resistance to rapamycin. In Botrytis cinerea, deletion of Bcfkbp12 caused a reduction of the virulence of strain T4, while knockout of the fkbp12 gene did not affect the pathogenic development of strain B05.10 [32]. These findings demonstrated that the fkbp12 homologous genes exhibited conservation in terms of rapamycin resistance. However, there is significant variation in other biological functions among fungal species, which may be attributed to species or strain evolutionary divergence.

Revised version: Our findings indicated that the TOR signaling pathway is pivotal in A. flavus, we have observed that rapamycin has a significant impact on the growth and sporulation of A. flavus (Fig. 1A). Additionally, it exhibits a strong inhibitory effect on sclerotia production and aflatoxin biosynthesis (Fig. 1B, 1C). The inhibitory effect has been documented in various fungal species, including Candida albicans [51], Cryptococcus neoformans [51], Podospora anserine [27], A. nidulans [28], Magnaporthe oryzae [39], and F. graminearum [23]. This suggested that the TOR signaling pathway may have a conserved function in filamentous fungi. The number of FK506-binding proteins varies among different species: there are four Fkbps proteins in S. cerevisiae and three Fkbps proteins in S. pombe [52]. In S. cerevisiae, rapamycin can bind to Fkbp protein, resulting in the irreversible inhibition of the G1 phase of the cell cycle and the regulation of cell growth [53]. In A. nidulans, the fkbp12 gene homologue fprA displays considerable sequence similarity with homologues identified in the S. cerevisiae [28]. In this study, based on the homologous alignment of the protein sequence of FprA, it was observed that four fkbp genes are present in A. flavus. Among these, deletion of fkbp3 exhibited enhanced tolerance to rapamycin (Fig. 2A). So we speculated that rapamycin forms a complex with Fkbp3 in A. flavus, subsequently targeting downstream target genes to exert its biological function. We found that Fkbp3 is involved in sclerotia formation and aflatoxin biosynthesis in A. flavus (Fig. 2A, S3), and subsequent analysis revealed that the underlying mechanism may be associated with succinylation modification. The succinylation sites of Fkbp3 were identified through the utilization of succinylation proteomics data. By constructing point mutant strains and conducting phenotype experiments, we found that the succinylation of Fkbp3 at residue K19 is involved in rapamycin resistance and aflatoxin biosynthesis (Fig. 2E, 2G). Our results provided new insights to the understanding of post-translational modifications in the regulation of TOR pathway components. Previous studies have demonstrated that Fkbp12 has been identified as the receptor for FK506 and rapamycin, which inhibits calcineurin and target rapamycin complex 1 (TORC1), respectively [53]. The deletion of Fkbp12 confers the resistance to rapamycin in C. albicans [51] and C. neoformans [51]. The entomopathogenic fungus Beauveria bassiana contains three putative fkbp genes, and disruption of Bbfkbp12 significantly increases the resistance to rapamycin and FK506 [54]. The plant pathogenic fungus F. graminearum harbors three Fkbps proteins, among which deletion of Fgfkbp12 leads to resistance to rapamycin [23]. Thus, we propose that these Fkbps may exhibit analogous functionalities in conferring resistance to rapamycin. In Botrytis cinerea, deletion of Bcfkbp12 caused a reduction of the virulence of strain T4, while knockout of the fkbp12 gene did not affect the pathogenic development of strain B05.10 [32]. These findings demonstrated that the fkbp12 homologous genes exhibited conservation in terms of rapamycin resistance. However, there is significant variation in other biological functions among fungal species, which may be attributed to species or strain evolutionary divergence. Our study indicated the pivotal role of Fkbp3 in modulating the effects of rapamycin on A. flavus, Fkbp3 is also crucial for aflatoxin biosynthesis. Targeted inhibition of Fkbp3 could potentially suppress aflatoxin production, given its significance, Fkbp3 can be developed as a drug target to reduce aflatoxin biosynthesis, thereby reducing the risk of aflatoxin contamination in agricultural products.

Earlier version: The Tor kinase functions as a fundamental component of the TOR signaling pathway, which exhibits evolutionary conservation across fungi, plants, and mammals [52]. In S. cerevisiae, the Tor kinase can interact with a diverse range of proteins, resulting in the formation of complex structures known as TORC1 and TORC2 [53]. TORC1 can detect various signals, including nutrients, growth factors, and environmental stress. It plays a crucial role in regulating cellular processes such as gene transcription, protein translation, ribosome synthesis, and autophagy [54]. A characteristic that distinguishes TORC1 is its interaction with KOG1 (Raptor in mammals) and its sensitivity to the rapamycin-FKBP12 complex [55]. Almost all other eukaryotes, including plants, worms, flies, and mammals, have a single tor gene, whereas S. cerevisiae has two tor genes [54]. The Tor kinase is widely conserved both in terms of its structural characteristics and its role as the target of Fkbp-rapamycin. Typically, the Tor kinase exhibits similar structural characteristics, including the HEAT-like, FAT, FRB, kinase, and FATC domains. The FRB domain is the binding region of the Fkbp-rapamycin complex, and rapamycin resistance is usually caused by the disruption of the FRB domain [54]. The comparison of homology and the analysis of phylogenetics revealed that there is only one predicted torA gene in A. flavus. To ascertain its functionality, several ectopic transformants were obtained through multiple protoplast transformations, however, no successful knock-out transformants were obtained. We speculated that the deletion of the torA gene may be lethal in A. flavus. This observation aligns with the situation observed in F. graminearum [23]. Previous findings demonstrated that the Tor kinase could function as a molecular switch connecting an iron cue to defend against pathogen infection in C. elegans [56]. In A. fumigatus, the Tor kinase represents a central regulatory node controlling genes and proteins involved in nutrient sensing, stress response, cell cycle progression, protein biosynthesis, and degradation, as well as adaptation to low iron conditions[57]. In summary, the Tor kinase acts as a central regulator in the TOR signaling pathway, playing a crucial role in the regulation of the cell cycle, protein synthesis, and cellular energy metabolism.

Revised version: The Tor kinase functions as a fundamental component of the TOR signaling pathway, which exhibits evolutionary conservation across fungi, plants, and mammals [55]. In S. cerevisiae, the Tor kinase can interact with a diverse range of proteins, resulting in the formation of complex structures known as TORC1 and TORC2 [44]. TORC1 can detect various signals, including nutrients, growth factors, and environmental stress. It plays a crucial role in regulating cellular processes such as gene transcription, protein translation, ribosome synthesis, and autophagy [45]. A characteristic that distinguishes TORC1 is its interaction with KOG1 (Raptor in mammals) and its sensitivity to the rapamycin-FKBP12 complex [56]. The Tor kinase is widely conserved both in terms of its structural characteristics and its role as the target of Fkbp-rapamycin [57]. Typically, the Tor kinase exhibits similar structural characteristics, including the HEAT-like, FAT, FRB, kinase, and FATC domains [57]. The FRB domain is the binding region of the Fkbp-rapamycin complex, and rapamycin resistance is usually caused by the disruption of the FRB domain [45]. Homology analysis and phylogenetic studies have identified a single predicted torA gene within the A. flavus genome. To elucidate the functional role of torA in A. flavus, our initial approach involved an attempt to generate an torA deletion strain through homologous recombination. However, these efforts were unsuccessful. We speculated that the deletion of the torA gene may be lethal in A. flavus. This observation aligns with the situation observed in F. graminearum [23]. Therefore, to investigate the function of TorA in A. flavus, we constructed a mutant strain (xylPtorA) using an xylose-inducible promoter, which allows for conditional induction with the addition of xylose. We observed that the xylPtorA mutant is completely unable to grow in YGT medium without xylose, whereas partial growth is restored in YXT medium with the addition of xylose (Fig. 3A). These results indicated that TorA is essential for the growth of A. flavus. We also discovered that the Tor kinase is involved in the vegetative growth, asexual development and aflatoxin production of A. flavus (Fig. 3A, 3B). In addition, the the xylPtorA mutant shows a high sensitivity to rapamycin and SDS (Fig. 3A). Our findings have demonstrated that the TorA kinase plays a crucial role in the processes of morphogenesis, aflatoxin biosynthesis, and multiple stress responses in A. flavus. Previous findings demonstrated that the Tor kinase could function as a molecular switch connecting an iron cue to defend against pathogen infection in C. elegans [58]. In A. fumigatus, the Tor kinase represents a central regulatory node controlling genes and proteins involved in nutrient sensing, stress response, cell cycle progression, protein biosynthesis, and degradation, as well as adaptation to low iron conditions[59]. In summary, the Tor kinase acts as a central regulator in the TOR signaling pathway, playing a crucial role in the regulation of the cell cycle, protein synthesis, and cellular energy metabolism. Its influence on key cellular processes makes it a promising target for future research to control fungal growth and aflatoxin production.

Earlier version: In S. cerevisiae, the protein kinase Sch9 is a major downstream effector of TORC1, and this kinase is known to have crucial functions in stress resistance, longevity, and nutrient sensing [44]. In addition, the TORC1-Sch9 pathway acts as a crucial mediator of chronological lifespan [44]. In filamentous fungi, an intricate network of interactions exists between the TOR and HOG pathways. In F. graminearum [38], FgSch9 and FgHog1 mutants displayed heightened susceptibility to osmotic and oxidative stresses. The phosphorylation level of FgHog1 increased significantly in FgSch9 mutant. Additionally, the affinity capture-MS assay showed that the Tor kinase was among the proteins co-purifying with FgSch9. These results suggest that FgSch9 serves as an intermediary for the TOR and HOG signaling pathways [38]. Furthermore, FgSch9 plays a crucial role in the regulation of various biological processes such as hyphal differentiation, asexual development, virulence, and DON biosynthesis[38]. In A. fumigatus, the SchA mutant was sensitive to rapamycin, high concentrations of calcium and hyperosmotic stress, and SchA was involved in iron metabolism. The SchA mutant exhibited enhanced phosphorylation of SakA in response to osmotic stress [29]. In this study, we discovered that Sch9 is involved in regulating aflatoxin biosynthesis (Fig. 3C), and responses to osmotic stress and rapamycin stress (Fig. 3B, 3D). The S_TKc domain deletion strain (sch9ΔS_TKc), S_TK_X domain deletion strain (sch9ΔS_TK_X), and sch9K340A mutant displayed a similar phenotype to the Δsch9 strain (Fig. 3B, 3C, 3D). These findings revealed that Sch9 plays a significant role in the biological synthesis of aflatoxin and in responding to various stresses. This role is primarily mediated through its kinase domain, kinase-C domain, and ATP binding site K340. We also found that the phosphorylation level of Hog1 increased significantly in the Δsch9 strain under osmotic stress (Fig. 3E). Based on the above results, we hypothesized that Sch9 was involved in the osmotic stress response by modulating the HOG pathway.

Revised version: The serine/threonine protein kinase Sch9, a critical downstream effector of TORC1 in S. cerevisiae, has been extensively studied for its roles in stress resistance, longevity, and nutrient sensing [46]. In addition, the TORC1-Sch9 pathway is also recognized as crucial mediator of chronological lifespan [46]. In filamentous fungi, an intricate network of interactions exists between the TOR and HOG pathways. In F. graminearum [38], FgSch9 and FgHog1 mutants displayed increased susceptibility to osmotic and oxidative stresses. The phosphorylation level of FgHog1 increased significantly in FgSch9 mutant. Additionally, the affinity capture-MS assay showed that the Tor kinase was among the proteins co-purifying with FgSch9. These results suggested that FgSch9 serves as an intermediary for the TOR and HOG signaling pathways [38]. Furthermore, FgSch9 plays a crucial role in the regulation of various biological processes such as hyphal differentiation, asexual development, virulence, and DON biosynthesis[38]. In A. fumigatus, the SchA mutant was sensitive to rapamycin, high concentrations of calcium and hyperosmotic stress, and SchA was involved in iron metabolism. The SchA mutant exhibited enhanced phosphorylation of SakA in response to osmotic stress [29]. In this study, we discovered that Sch9 is involved in regulating aflatoxin biosynthesis (Fig. 4C), and responses to osmotic and rapamycin stress (Fig. 4B, 4D). The S_TKc domain deletion strain (sch9ΔS_TKc), S_TK_X domain deletion strain (sch9ΔS_TK_X), and sch9K340A mutant displayed a similar phenotype to the Δsch9 strain (Fig. 4B, 4C, 4D). These findings revealed that Sch9 plays a significant role in the biological synthesis of aflatoxin and in responding to various stresses. This role is primarily mediated through its kinase domain, kinase-C domain, and ATP binding site K340. We also found that the phosphorylation level of Hog1 increased significantly in the Δsch9 strain under osmotic stress (Fig. 4E). Based on the above results, we hypothesized that Sch9 was involved in the osmotic stress response by modulating the HOG pathway in A. flavus.

Earlier version: Tap42, a protein that associates with phosphatase 2A, serves as the direct target of the Tor kinase in yeast. Tap42 plays a crucial role in various Tor functions, particularly in the regulation of transcriptional processes [58]. The genome of A. flavus contains a homolog of tap42 named tapA (AFLA_092770), which is predicted to encode a 355 amino-acid protein (sharing 31.9% identity with S. cerevisiae Tap42). To further determine the function of the TapA regulator in the TOR signaling pathway of A. flavus, we conducted targeted gene deletion experiments on tapA. We failed to obtain a null mutant, we hypothesised that tapA knockout was lethal in A. flavus. In S. pombe, the protein Tip41, which interacts with Tap42, plays a significant role in the cellular responses to nitrogen sources by regulating type 2A phosphatases [59]. In S. cerevisiae, Tap42-phosphatase complexes associate with TORC1, whereas Tip41 functions to attenuate TORC1 signaling by stimulating the PP2A phosphatase [45]. Unlike its yeast counterpart Tip41, TIPRL, a mammalian Tip41-like protein, exerts a positive influence on mTORC1 signaling by interacting with PP2Ac [60]. While the interaction between Tip41 and Tap42 was not observed in the filamentous fungi. The yeast two-hybrid (Y2H) assay showed that FgTip41 interacts with the phosphatase FgPpg1 rather than FgTap42, while the FgTap42 bind with phosphatases FgPp2A, implying that FgTap42, FgPpg1, and FgTip41 form a heterotrimer in F. graminearum [23]. In the rice blast fungus M. oryzae, MoTip41 does not interact with MoTap42. However, MoTip41 can bind with the phosphatase MoPpe1, thereby facilitating crosstalk between the TOR pathway and the cell wall integrity pathway [40]. In this study, we found that TipA played a crucial role in sclerotia production and cell wall stress in A. flavus (Fig. 4B, 4H). However, the interaction between TapA and TipA is not yet clear in A. flavus.

Revised version: The phosphorylation state of Tap42 is crucial for its interaction with PP2A subunits and is regulated by a balance between Tor-mediated phosphorylation and Cdc55/Tpd3-mediated dephosphorylation in yeast [60]. The genome of A. flavus contains a homolog of tap42 named tapA (AFLA_092770), which is predicted to encode a 355 amino-acid protein (sharing 31.9% identity with S. cerevisiae Tap42). To further determine the function of the TapA regulator in the TOR signaling pathway of A. flavus, we conducted targeted gene deletion experiments on tapA. We failed to obtain a null mutant, we hypothesised that tapA knockout was lethal in A. flavus. To investigate the function of TapA, we employed homologous recombination, to replace the native promoter with the strong gpdA promoter, thereby constructing an overexpression strain (OE::tapA) with the enhanced promoter. Compared with WT strain, OE::tapA strain were more sensitive to rapamycin and SDS stress, indicating that tapA gene plays an important role in regulating cell wall stress (Fig. 5B, 5F). In S. pombe, the protein Tip41, which interacts with Tap42, plays a significant role in the cellular responses to nitrogen sources by regulating type 2A phosphatases [61]. In S. cerevisiae, Tap42-phosphatase complexes associate with TORC1, whereas Tip41 functions to attenuate TORC1 signaling by stimulating the PP2A phosphatase [48]. Unlike its yeast counterpart Tip41, TIPRL, a mammalian Tip41-like protein, exerts a positive influence on mTORC1 signaling by interacting with PP2Ac [62]. While some functions of Tap42 and its binding partners Tip41 may be conserved, the interaction between Tip41 and Tap42 was not observed in the filamentous fungi. The yeast two-hybrid (Y2H) assay showed that FgTip41 interacts with the phosphatase FgPpg1 rather than FgTap42, while the FgTap42 bind with phosphatases FgPp2A, implying that FgTap42, FgPpg1, and FgTip41 form a heterotrimer in F. graminearum [23]. In the rice blast fungus M. oryzae, MoTip41 does not interact with MoTap42. However, MoTip41 can bind with the phosphatase MoPpe1, thereby facilitating crosstalk between the TOR pathway and the cell wall integrity pathway [40]. In A. flavus, the OE::tapA and ΔtipA strains showed increased sensitivity to cell wall stress compared with the WT strain(Fig. 5B, 6H). Based on the conserved roles of Tap42 and Tip41 in yeast and other filamentous fungi, we hypothesize that TapA and TipA in A. flavus also participate in cell wall integrity pathway. In summary, our study provided important insights into the role of TapA and TipA in the TOR signaling pathway of A. flavus, particularly in relation to cell wall stress. However, their interaction and implications for TOR and CWI signaling in A. flavus remain to be elucidated.

Earlier version: In S. cerevisiae, Tap42 can bind and regulate various PP2A phosphatases, including Pph3, Pph21, Pph22, Sit4, and Ppg1. This interaction plays a crucial role in the TOR signaling pathway, particularly in response to rapamycin treatment or nutrient deprivation [45]. The phosphatase Sit4 plays a crucial role in regulating cell growth, morphogenesis, and virulence in C. albicans [61]. FgSit4 and FgPpg1 are associated with various cellular processes, including mycelial growth, conidiation, DON biosynthesis, and virulence in F. graminearum [23]. In addition, FgSit4 and FgPpg1 are involved in the regulation of various signaling pathways, such as the CWI pathway and lipid droplets biogenesis regulation pathway [23]. In this study, we found that the phosphatase SitA and Ppg1 play critical role in vegetative growth, conidiation, sclerotia formation, and aflatoxin biosynthesis in A. flavus (Fig. 5A, 5B, 5C). Additionally, the phosphatase SitA and Ppg1 also affect pathogenicity and LD biogenesis (Fig. 6A, 8C, 8D). In addition, SitA, Ppg1, and TipA exhibits similar functions in sclerotia production and cell wall stress response. Based on these observations, we hypothesized that TapA-phosphatase complexes play a collaborative role in the regulation of the sclerotia formation and cell wall stress response in A. flavus.

Revised version: In S. cerevisiae, Tap42 can bind and regulate various PP2A phosphatases, including Pph3, Pph21, Pph22, Sit4, and Ppg1. This interaction plays a crucial role in the TOR signaling pathway, particularly in response to rapamycin treatment or nutrient deprivation [48]. These phosphatases play important roles in other fungi and are involved in a variety of cellular processes, including growth, morphogenesis, and virulence. For example, the phosphatase Sit4 plays a crucial role in regulating cell growth, morphogenesis, and virulence in C. albicans [63]. Similarly, FgSit4 and FgPpg1 are associated with various cellular processes, including mycelial growth, conidiation, DON biosynthesis, and virulence in F. graminearum [23]. These findings indicated the conserved role of PP2A phosphatases in fungal biology. In addition to their roles in growth and development, FgSit4 and FgPpg1 are involved in the regulation of various signaling pathways, such as the CWI pathway and lipid droplets biogenesis regulation pathway [23]. In this study, we found that the phosphatase SitA and Ppg1 play critical role in vegetative growth, conidiation, sclerotia formation, and aflatoxin biosynthesis in A. flavus (Fig. 7A, 7B, 7C). Additionally, the phosphatase SitA and Ppg1 also affect pathogenicity and LD biogenesis (Fig. 8A, 10C, 10D). This highlights the multifaceted influence of PP2A phosphatases on cellular signaling.

Earlier version: After the deletion of the sitA and ppg1 genes, the expression level of genes related to cell wall synthesis has decreased significantly, resulting in damage to the cell wall and showing a higher sensitivity to cell wall integrity stress (Fig. 7C). Western blotting results showed that the phosphorylation level of MAPK kinase Slt2 has increased significantly in the ΔsitA and Δppg1 strains under cell wall stress, indicating that SitA and Ppg1 exert a negative regulatory effect on the CWI pathway (Fig. 7D). In M. oryzae, MoPpe1 and MoSap1 serve as an adaptor complex that connects the CWI and TOR signaling pathways. Activation of the TOR pathway results in the inhibition of the CWI pathway [39]. In addition, the phosphatase MoPpe1 interacts with the phosphatase-associated protein MoTip41, facilitating the exchange of information between the TOR and CWI signaling pathways [39]. These results suggest that the phosphatase Ppe1 engages in interactions with various proteins to effectively coordinate the TOR and CWI signaling pathways, thereby regulating the growth and pathogenicity of the rice blast fungus M. oryzae [39]. Taken together, we speculated that the phosphatase SitA and Ppg1 might exhibit differential regulation in the CWI and TOR pathways across species in order to fulfill their respective functions.

Revised version: After the deletion of the sitA and ppg1 genes, the expression level of genes related to cell wall synthesis has decreased significantly, resulting in damage to the cell wall and showing a higher sensitivity to cell wall integrity stress (Fig. 9C). Western blotting results showed that the phosphorylation level of MAPK kinase Slt2 has increased significantly in the ΔsitA and Δppg1 strains under cell wall stress, indicating that SitA and Ppg1 exert a negative regulatory effect on the CWI pathway (Fig. 9D). In M. oryzae, MoPpe1 and MoSap1 serve as an adaptor complex that connects the CWI and TOR signaling pathways. The activation of the TOR pathway leads to the inhibition of the CWI pathway, highlighting a critical regulatory interplay between these two pathways in fungal biology [39]. We speculated that SitA and Ppg1 might exhibit similar regulation within the CWI and TOR pathways across different fungal species. In addition, the phosphatase MoPpe1 interacts with the phosphatase-associated protein MoTip41, facilitating the exchange of information between the TOR and CWI signaling pathways [39]. These results suggested that the phosphatase Ppe1 engages in interactions with various proteins to effectively coordinate the TOR and CWI signaling pathways, thereby regulating the growth and pathogenicity of the rice blast fungus M. oryzae [39]. Given the conserved functions of PP2A phosphatases and the observed roles of SitA, Ppg1, TapA, and TipA in cell wall stress response, we hypothesized that TapA-phosphatase complexes may collaborate to regulate cell wall integrity processes in A. flavus. This hypothesis is supported by the similar functions exhibited by these proteins.

Earlier version: In F. graminearum, the phosphatase complex FgNem1/Spo7-FgPah1 cascade plays significant roles in fungal development, DON production, and lipid droplet (LD) biogenesis [24]. In this study, our results indicated that Nem1 and Spo7 are important for vegetative growth and secondary metabolism of A. flavus (Fig. 9A, 9B). We have also discovered that LD biogenesis is facilitated by the Nem1/Spo7 and SitA/Ppg1 pathways (Fig. 8C, 8D, 9E). These observations suggest that the PP2A phosphatases play a central role as core phosphatases in the LD biogenesis signaling network. However, there are likely other kinases involved in the regulation of LD biogenesis. To gain a comprehensive understanding of the precise molecular mechanism, further research should focus on investigating the interplay between phosphatases and other kinases.

Revised version: In F. graminearum, the phosphatase complex FgNem1/Spo7-FgPah1 cascade plays significant roles in fungal development, DON production, and lipid droplet (LD) biogenesis [24]. Our study explored the conservation and functional significance of Nem1 and Spo7 homologs in A. flavus, revealing their pivotal roles in vegetative growth and secondary metabolism (Fig. 11A, 11B). Our results indicated that lipid droplet biogenesis is facilitated by the Nem1/Spo7 and SitA/Ppg1 pathways in A. flavus (Fig. 10C, 10D, 10E). These observations suggested a conserved and central role for PP2A phosphatases as core regulators within the LD biogenesis signaling network, highlighting their importance in cellular lipid metabolism and energy storage. While our findings underscored the significance of phosphatases in LD biogenesis, it is also plausible that other kinases contribute to the regulation of this process. To gain a comprehensive understanding of the precise molecular mechanism, it is essential to investigate the interplay between phosphatases and other kinases.

Earlier version: In conclusion, we identified the complex regulatory network that provides insight into the interplay between the TOR, HOG, and CWI signaling pathways, and how they collectively regulate the development, aflatoxin biosynthesis, and pathogenicity of A. flavus. By elucidating the functions of genes associated with the TOR signaling pathway, our research could establish a solid foundation for comprehending the molecular mechanism underlying the aflatoxin biosynthesis of A. flavus. This, in turn, can provide novel insights into the management of aflatoxin contamination and the mitigation of A. flavus infection.

Revised version: In conclusion, we identified the complex regulatory network that provides important insight into the interplay between the TOR, HOG, and CWI signaling pathways, and how they collectively regulate the development, aflatoxin biosynthesis, and pathogenicity of A. flavus. By elucidating the functions of genes associated with the TOR signaling pathway, our research could establish a solid foundation for comprehending the molecular mechanism underlying the aflatoxin biosynthesis of A. flavus. Our results could also contribute to the development of targeted strategies to manage aflatoxin contamination and mitigate the impact of A. flavus infections.

Materials and methods

Strains and cultural conditions

Earlier version: The A. flavus strains in this study were shown in Table S1. All strains were cultivated on three different types of agar media: potato dextrose agar (PDA), yeast extract-sucrose agar (YES), and yeast extract-glucose agar (YGT). The growth and conidiation assays were conducted at 37℃. Additionally, aflatoxin production was assessed using YES liquid medium at 29℃.

Revised version: The A. flavus strains in this study were shown in Table S1. All strains were cultivated on four different types of agar media: potato dextrose agar (PDA), yeast extract-sucrose agar (YES), yeast extract-glucose agar (YGT) and yeast extract-xylose agar (YXT). The growth and conidiation assays were conducted at 37℃. Additionally, aflatoxin production was assessed using YES and YXT liquid medium at 29℃.

Constructions of gene deletion and complemented mutants

Earlier version: The construction of gene deletion and complemented strains was carried out using the SOE-PCR strategy, following previously established protocols [62]. The primers utilized for the amplification of the flanking sequences of each gene can be found in Table S2. The putative gene deletion mutants were validated through PCR assays using the appropriate primers (Fig. S2), and their confirmation was subsequently obtained through sequencing analysis.

Revised version: The construction of gene deletion and complemented strains was carried out using the SOE-PCR strategy, following previously established protocols [64]. To generate the xylPtorA mutant of A. flavus, an torA xylose promoter mutant cassette was synthesized via overlapping extension PCR. This cassette was designed to incorporate an upstream fragment of the torA gene, a selectable marker gene (pyrG from Aspergillus fumigatus), the xyIP xylose-inducible promoter [65], and a fragment of the torA coding sequence. Subsequently, this composite cassette was integrated into the A. flavus genome via homologous recombination, specifically replacing the native torA promoter with the inducible xylose promoter. To construct over-expression cassettes for the tapA gene, we utilized a molecular cloning strategy that entailed the fusion of PCR-amplified fragments from the pyrG-selectable marker and the gpdA promoter. Subsequently, employing a homologous recombination approach, we inserted this composite fragment upstream of the tapA gene, effectively replacing its native promoter with the stronger gpdA promoter. The primers utilized for the amplification of the flanking sequences of each gene can be found in Table S2. The putative gene deletion mutants were validated through PCR assays using the appropriate primers (Fig. S2), and their confirmation was subsequently obtained through sequencing analysis.

Growth, conidiation, and sclerotia analysis

Earlier version: To evaluate mycelial growth and conidia formation, 1 μL of 106 spores/mL conidia of all strains was point inoculated onto the center of the PDA medium, and then cultured at 37℃ for 5 days in the dark. The conidia were washed with 0.07% Tween-20 and counted using a hemocytometer and microscope [43]. For sclerotia formation analysis, each strain was inoculated and cultivated on CM agar medium at 37℃ in the dark for 7 days, and then 75% ethanol was used to wash away mycelia and conidia on the surface of the medium [63]. These assays were repeated at least three times.

Revised version: To evaluate mycelial growth and conidia formation, 1 μL of 106 spores/mL conidia of all strains was point inoculated onto the center of the PDA medium, and then cultured at 37℃ for 5 days in the dark. The conidia were washed with 0.07% Tween-20 and counted using a hemocytometer and microscope [43]. For sclerotia formation analysis, each strain was inoculated and cultivated on CM and YXT agar medium at 37℃ in the dark for 7 days, and then 75% ethanol was used to wash away mycelia and conidia on the surface of the medium [66]. These assays were repeated at least three times.

Extraction and quantitative analysis of aflatoxin

Earlier version: To extract and quantify aflatoxins (AFs) production, 1 μL of 106 spores/mL conidia of all strains was point inoculated onto liquid YES medium, and cultures were then incubated at 29℃ in the dark for 6 days. AFs were extracted from the media and detected using thin layer chromatography (TLC) [63]. In addition, the standard AFB1 was employed for visual comparison, and AFs were quantitatively analyzed using the Image J tool. The assays were repeated at least three times.

Revised version: To extract and quantify aflatoxins (AFs) production, 1 μL of 106 spores/mL conidia of all strains was point inoculated onto liquid YES medium, and cultures were then incubated at 29℃ in the dark for 6 days. AFs were extracted from the media and detected using thin layer chromatography (TLC) [66]. In addition, the standard AFB1 was employed for visual comparison, and AFs were quantitatively analyzed using the Image J tool. The assays were repeated at least three times.

Stress assay

Earlier version: To determine the roles of the gene in A. flavus responsing to various environmental stresses, 1 μL of 106 spores/mL conidia of all strains were point inoculated onto these PDA solid plates supplemented with various agents. These agents included 1M NaCl, 1M KCl, and 1M CaCl2 for inducing osmotic stress, as well as 100 µg/mL SDS, 200 µg/mL Calcofluor white (CFW), and 200 µg/mL Congo red (CR) for inducing cell wall stress. After 5 days of incubation at 37℃ in the dark, the relative rates of inhibition were determined [43]. The assays were repeated at least three times.

Revised version: To determine the roles of the gene in A. flavus responsing to various environmental stresses, 1 μL of 106 spores/mL conidia of all strains were point inoculated onto these PDA solid plates supplemented with various agents. These agents included 1M NaCl, 1M KCl, and 1M CaCl2 for inducing osmotic stress, as well as 100 µg/mL or 300 µg/mL SDS, 200 µg/mL Calcofluor white (CFW), and 200 µg/mL Congo red (CR) for inducing cell wall stress. After 5 days of incubation at 37℃ in the dark, the relative rates of inhibition were determined [67]. The assays were repeated at least three times.

Seeds infection assays

Earlier version: To measure the pathogenicity of A. flavus, all peanut and maize seeds were washed in 0.05% sodium hypochlorite, 75% ethanol, and sterile water, respectively. Subsequently, all peanut and maize seeds were immersed in sterile water containing 106 spores/mL of all strains at 29℃ in the dark for 60 min. After the conidia adhesion, the seeds were cultured at 29℃ for 6 days in the dark [64].

Revised version: To measure the pathogenicity of A. flavus, all peanut and maize seeds were washed in 0.05% sodium hypochlorite, 75% ethanol, and sterile water, respectively. Subsequently, all peanut and maize seeds were immersed in sterile water containing 106 spores/mL of all strains at 29℃ in the dark for 60 min. After the conidia adhesion, the seeds were cultured at 29℃ for 6 days in the dark [67].

Quantitative RT-PCR

Earlier version: For the real-time quantitative PCR (qPCR) assays, the A. flavus strains were cultivated in the PDA medium at 29℃ for 3 days in the dark. Total RNA was extracted from the mycelia of all strains, and cDNA was synthesized with the First Strand cDNA Synthesis Kit, and then the cDNA was used as the template for qRT-PCR analysis with SYBR Green qPCR mix. The relative transcript level of each gene was calculated by the 2-ΔΔCTmethod [65].

Revised version: For the real-time quantitative PCR (qPCR) assays, the A. flavus strains were cultivated in the PDA medium at 29℃ for 3 days in the dark. Total RNA was extracted from the mycelia of all strains, and cDNA was synthesized with the First Strand cDNA Synthesis Kit, and then the cDNA was used as the template for qRT-PCR analysis with SYBR Green qPCR mix. The relative transcript level of each gene was calculated by the 2-ΔΔCTmethod [68].

Western blot analysis

To extract the total protein, the mycelia of each strain were cultivated in YES liquid medium for 48 h, then the sample was grinded under liquid nitrogen to a fine powder and transferred to a 1.5 mL microcentrifuge tube added with RIPA lysis buffer and 1 mM PMSF protease inhibitor. Proteins were separated by 12% SDS-PAGE and transferred onto a polyvinylidene fluoride (PVDF) membrane. Phosphorylation levels of Hog1 and Slt2 were detected by Phospho-p38 MAPK antibody (Cell Signaling Technology, Boston, MA, USA) and Phospho-p44/42 MAPK (Erk 1/2) antibody (Cell signaling Technology, MA, USA), respectively. The experiment was conducted three times independently.

Statistical analysis

All data were measured with means ± standard deviation (SD) of three biological replicates. Graph Pad Prism 7 software was used for conducting statistical and significance analysis, and p-values < 0.05 represented statistical significance. Student’s t-test was used to compare the two means for differences, whereas Tukey’s multiple comparisons were used for testing multiple comparisons.

Supporting information